Two-way Dispatched function in Sonic hedgehog shedding and transfer to high-density lipoproteins

Author Response

The following is the authors’ response to the original reviews.

Reviewer 1 (public):

1. “It is unclear whether new in vivo experiments were conducted for this study”.

All in vivo experiments were conducted for this study by using previously published fly stocks to directly compare N- and C-terminal shedding side-by-side in two Hh-dependent developmental systems. This is now clearly stated in the revised supplement (Fig. S8). We also conducted these experiments because previous in vivo studies in flies often relied on Hh overexpression in the fat body, raising questions about their physiological relevance. Our in vivo analyses of Hh function in wing and eye discs are more physiologically relevant and can explain the previously reported presence of non-lipidated bioactive Hh in disc tissue (PMID: 23554573).

1. “A critical shortcoming of the study is that experiments showing Shh secretion/export do not include a Shh(-) control condition. Without demonstration that the bands analyzed are specific for Shh(+) conditions, these experiments cannot be appropriately evaluated”.

The Cell Signaling Technology C9C5 anti-Shh antibody used in our study is highly specific against Shh, and it has been used in over 60 publications. C9C5 even lacks cross-reactivity with highly similar Ihh or Dhh (https://www.cellsignal.com/products/primary-antibodies/shh-c9c5-rabbit-mab/2207?_requestid=1528451). We confirmed C9C5 specificity repeatedly (one example is shown below; another quality control that includes media of mock-transfected cells is now shown in Fig. S1) and never observed unspecific bands under any experimental condition. As shown below, C9C5 and R&D AF464 anti-Shh antibodies (the latter were previously used in our lab) detect the same bands.

Shh immunoblot.

R&D 8908-SH served as a size control for full-length dual-lipidated Shh, and C25S;26-35Shh served as a size control for N-terminally truncated monolipidated Shh. Both C25SShh bands are specific: One represents the full-length protein and the bottom band represents N-truncated processed proteins. The blot was first incubated with antibody AF464 and reincubated (after stripping) with the much more sensitive antibody C9C5.

1. “A stably expressing Shh/Hhat cell line would reduce condition to condition and experiment to experiment variability”.

We agree and therefore have previously aimed to establish stable Hhat-expressing cell lines. However, we found that long-term Hhat overexpression eliminated transfected cells after several passages, or cells gradually ceased to express Hhat. This prevented us from establishing stable cell lines co-expressing Shh/Hhat despite several attempts and different strategies. Instead, we established transient co-expression of Shh/Hhat from the same mRNA as the next-best strategy for reliable near-quantitative Shh palmitoylation in our assays.

1. “Unusual normalization strategies are used for many experiments, and quantification/statistical analyses are missing for several experiments”.

We repeated all qPCR assays to eliminate this shortcoming. Biological activities and transcriptional responses of palmitoylated Shh and non-palmitoylated C25AShh are now directly compared and quantified (revised Fig. 4A,B, newly included Fig. 6, revised Fig. S5B). The original comparison of both proteins with dual-lipidated R&D 8908-SH is still important in order to show that both Shh and C25AShh in serum-containing media have equally high, and not equally low, activities because R&D 8908-SH is generally seen as the Shh form with the highest biological activity. These comparisons are therefore still discussed in the main manuscript text and are now shown in Fig. S5E.

1. “The study provides a modest advance in the understanding of the complex issue of Shh membrane extraction”

We believe that the revised manuscript advances our understanding of Shh membrane extraction beyond the modest in three important ways. First, although Disp was indeed known as a furin-activated Hh exporter, our findings show for the first time that furin activation of Disp is strictly linked to proteolytic Shh processing as the underlying release mode, fully consistent with data obtained from the Disp-/- cells.

Second, Scube2 was known as a Shh release enhancer and several lipoproteins were previously shown to play a role in the process, but our findings are the first to show that synergistic Disp/Scube2 function depends on the presence of lipoprotein and that HDL (but no other lipoprotein) accepts free cholesterol or a novel monolipidated Shh variant from Disp. This challenges the dominant model of Scube2 chaperone function in Hh release and transport (PMID 22902404, PMID 22677548, PMID 36932157).

Third, we show that this Shh variant is fully bioactive, despite the lack of the palmitate. Therefore, N-palmitate is dispensable for Shh signaling to Ptch1 receptors, but only if the morphogen is released by, and physically linked to, HDL. In contrast, previously published studies analyzed monolipidated Shh variants in the absence of HDL, resulting in variably reduced bioactivity of these physiologically irrelevant forms. Therefore, our findings challenge the current dominating model of N-palmitate-dependent Shh signaling to Ptch1 (this model also does not postulate any role for lipoproteins, PMID 36932157) and essential roles of N-palmitate (stating that the N-palmitate is sufficient for signaling, PMID 27647915).

Reviewer 2 (public):

1. “However, the results concerning the roles of lipoproteins and Shh lipid modifications are largely confirmatory of previous results, and molecular identity/physiological relevance of the newly identified Shh variant remain unclear”.

We disagree with this assessment on several points. First, our findings do not confirm, but strongly challenge, the current dogma of Disp-mediated handover of dual-lipidated Shh to Scube2 as a soluble acceptor (instead of to HDL, PMID 36932157). Second, we report three new findings: Disp, Scube2, and lipoproteins all interact to specifically increase N-terminal Shh shedding, whereas C-terminal shedding is optional; Disp function depends on the presence of HDL; and HDL modulates Shh shedding (dual Shh shedding in the absence of HDL versus N-shedding and HDL association in its presence). Our work also directly determines the molecular identity of a previously unknown Shh variant as monolipidated (by RP-HPLC), HDL associated (by SEC and density gradient centrifugation), and fully bioactive (in two cell-based reporter assays).

Third, regarding the physiological relevance of our findings: Fig. S8 demonstrates that deletion of the N-terminal sheddase target site of Hh abolishes all Hh biofunction in Drosophila eye discs and wing discs, which strongly supports physiological relevance of N-terminal Hh shedding during release. N-terminal shedding is further consistent with in vivo findings of others. These studies showed that artificial monolipidated Shh variants (C25SShh and ShhN) generate highly variable loss-of-function phenotypes in vivo, but can also generate gain-of-function phenotypes if compared with the dual-lipidated cellular protein 1, 2, 3, 4, 5. These observations are difficult to align with the dominating model of essential N-palmitate function at the level of Ptch1 (PMID 36932157), because the lack of N-palmitate is expected to always diminish signaling in all tissue contexts and developmental stages. Our finding that dual-lipidated Shh is strictly released in a Disp/Scube2-controlled manner from producing cells, while artificial monolipidated Shh variants leak uncontrolled from the cellular surface, explains these seemingly paradoxical in vivo findings much better. This is because uncontrolled Shh release can increase Shh signaling locally (when physiological release would normally be prevented at this site 6 or time), while it can also decrease it (for example, in situations requiring timed pulses of Shh release and signaling 7, 8, 9, 10, 11). This is discussed in our manuscript (Discussion, first paragraph).

1. The molecular properties of the processed Shh variants are unclear – incorporation of cholesterol/palmitate and removal of peptides were not directly demonstrated…

We also disagree on this point. Our study is the only one that uses RP-HPLC and defined controls (dual-lipidated commercial R&D 9808-SH, dual-lipidated cellular proteins eluting at the same positions, non-lipidated or monolipidated controls, Fig. S1F-K) to compare the lipidation status of cellular and corresponding solubilized Shh and to determine their exact lipidation status (Figs. 1, 3, 5, Figs. S4, S6, S7). Co-expressed Hhat assures full Shh palmitoylation during biosynthesis (as shown in original Figs. 1A and S2F-K & S4A and as confirmed by R&D 9808-SH) as an essential prerequisite to reliably conduct and interpret these analyses. The removal of peptides is demonstrated by the increase in electrophoretic mobility of soluble forms, if compared with their dual-lipidated cellular precursor, because chemical delipidation results in a decrease in electrophoretic mobility in SDS-PAGE (as discussed in detail in 12 that we now cite in our work).

1. This (N-terminal palmitoylation status) is particularly relevant …, as the signaling activity of non-palmitoylated Hedgehog proteins is controversial.

We agree with this comment and are aware of the published data. However, in our work, we have demonstrated strong signaling activities by using C25AShh mutants that are fully impaired in their ability to undergo N-palmitoylation (Fig. 4, Fig. S5). These are highly bioactive if associated with HDL. Therefore, we do not see any ambiguity in our findings and suggest that the reports of others resulted from different experimental conditions.

1. A decrease in hydrophobicity is no proof for cleavage of palmitate, this could also be due to addition of a shorter acyl group.

As shown in the original manuscript, we have controlled for this possibility: RP-HPLC was established by using defined controls (dual-lipidated, non-lipidated, or monolipidated, Fig. S1F-K and corresponding color coding). Because the cellular Shh precursor prior to release was always dual-lipidated, whereas the soluble form was not, lipids were clearly lost during release (because a decrease in the hydrophobicity of soluble proteins is always shown relative to that in their dual-lipidated cellular precursors). The increase in electrophoretic mobility detected for the very same proteins in SDS-PAGE demonstrates delipidation during their release (please see my reply to point 2 above). Finally, the suggested possibility of palmitate exchange for shorter acyls during Shh release at the cell surface is extremely unlikely, as there is no known machinery to catalyze this exchange at the plasma membrane. Hh acylation only occurs in the ER membrane via Hhat 13.

1. “It would be important to demonstrate key findings in cells that secrete Shh endogenously”.

We now show that Panc1 cells release endogenous Shh in truncated form, as our transfected cells do (Fig. S1). Moreover, the experimental data shown in Fig. S8B demonstrate that engrailed-controlled expression of sheddase-resistant Hh variants in wing disc cells completely blocks endogenous Hh produced in the same cells by stalling Disp-mediated morphogen export. Both findings strongly support our key finding that N-processing is not optional but absolutely required to finalize Hh release.

1. Co-fractionation of Shh and ApoA1 is not convincing, as the two proteins peak at different molecular weights…. The authors could use an orthogonal approach, optimally a demonstration of physical interaction, or at least fractionation by a different parameter

Shifted Shh peaks upon physiologically relevant Shh transfer via Disp to HDL must be expected in SEC, because Shh association with HDL subfractions increases their size. Comparing relative peaks of Shh-loaded HDL with Shh-free reference HDL suggests 10-15 Shh molecules per HDL (adding 200kDa - 300kDa to its molecular mass). This is now stated in the revised manuscript (page 10, line 2).

Still, to further support direct Shh/HDL association, we analyzed high molecular weight Shh SEC fractions by subsequent RP-HPLC. This approach confirms direct physical interactions between cholesteroylated Shh and HDL (now shown in Fig. S6G).

We support this possibility further by density gradient centrifugation, again demonstrating that Shh and HDL interact physically (now shown in Fig. S6 E,F).

Recommendations from the reviewing editor:

1. “The authors should certainly tone down statements of novelty because much of the work is confirmatory in nature”

We followed this request in our revised manuscript and now clearly point out what was known and what we add to the concept of Disp and lipoprotein-mediated Hh export. Still, as outlined in our response to reviewer 2, our findings align with only one previously published model of lipoprotein-mediated Hh transport, while they do not support the most current models of Disp-mediated handover of dual-lipidated Shh to Scube2 (PMID 36932157) and essential signaling roles of N-palmitate at the level of the receptor Ptch1. Thus, our work should not be viewed solely as confirmatory of one of the many previous models, because at the same time it also contradicts the other models of Hh solubilization and transport.

1. “Inclusion of the Shh(-) control”

Please see our reply to reviewer 1 above. The Cell Signaling Technology C9C5 anti-Shh antibody used in our study is highly specific against Shh. We also carefully characterized the C9C5 antibody before any of the experiments shown in our work had been initiated. We never observed any unspecific C9C5 reactivity that otherwise would – of course – have prevented us from switching to this antibody from the AF464 antibodies that we had previously used. Consistent C9C5 antibody specificity is evident from the representative example shown below that was recently produced in our lab: no cellular proteins or TCA-precipitated serum-depleted media components from mock-transfected cells (left two lanes) react with C9C5.

Top left: C9C5 detects the cellular 45kDa Shh precursor and the 19 kDa signaling-active protein.

No unspecific signals are detected in untransfected cells and supernatants of such cells (left two lanes). Right: Loading control on the stripped blot.

1. “Clean up how the data are normalized for quantification”

Please see our reply to reviewer 1 above. Normalization has been changed for the indicated figures. We also repeated qPCR analyses and added new ones to the manuscript that include required controls. We also changed figure outlines in accordance with the request.

1. “The issue of a non-specific band of this Shh antibody is critical”

Please see our replies above. In our hands, unspecific C9C5 antibody binding was never observed.

1. “Regarding experimental rigor, I would add that the HPLC … should just show the real data points”

We agree and added individual data points to our revised manuscript.

Recommendations for the authors:

1. I would like to see the controls in the same figure with the experimental results.

We show antibody specificity controls together with released Shh in Fig. S1.

1. Figure 2 confirms previously published results. It was shown in PMC5811216 that Disp processing by furin is required for Shh release from producing cells.

Indeed, it was shown that furin processing of Disp increases Shh release (supposedly together with lipids), but we show here that furin-activated Disp specifically mediates proteolytic Shh shedding and loss of lipids – which is not the same. Indeed, we show this finding because we interpret it the other way around: Because it is known that furin activation of Disp increases Shh release by some means (PMC5811216), our observation that furin-mediated Disp activation specifically increases Shh shedding independently supports our model.

1. Figure 3: it is stated that there is no increase in Shh release into the media…

We removed this statement.

1. Figure S5: Scale bars are missing.

We added scale bars to the figures.

1. Figure 4: A direct comparison between wt Shh and C25A conditioned media for qPCR is needed.

We agree and repeated all experiments. Results confirm our previous findings and are shown in revised Fig. 4 and in Fig. S5.

1. What other components can be examined in addition to ApoA1 as a marker for HDL? Why is the Shh peak shifted to the left? What about exovesicles?

We also detected ApoE4, a mobile lipoprotein present on expanding (large) HDL (Figs. 5, 6, Figs S6, 7) 14. We also used density gradient centrifugation to support the Shh/HDL association. Regarding the leftwards Shh size shift relative to the major HDL peak in SEC, please refer to our explanation above – if loaded with Shh, a size increase of the respective HDL subfraction is expected. Finally, we did not test the role of exovesicles in our assays. However, due to their large size (60-120nm, HDL 7-12 nm), Shh associated with exovesicles should have eluted in the void volume of our gel filtration column. This we never observed.

1. Why is osteoblast differentiation used?

C3H10T1/2 osteoblast differentiation is strongly driven by Ihh and Shh activity and is established as a sensitive and robust assay. Still, following this reviewer’s advice, we conducted qPCR assays on these cells and in addition on NIH3T3 cells to support our findings.

Finally, we corrected all minor mistakes regarding spelling and figure labeling. We also improved the readability of the revised manuscript, as suggested by reviewer 2.

References

1. Gallet A, Ruel L, Staccini-Lavenant L, Therond PP. Cholesterol modification is necessary for controlled planar long-range activity of Hedgehog in Drosophila epithelia. Development 133, 407-418 (2006).

2. Porter JA, et al. Hedgehog patterning activity: role of a lipophilic modification mediated by the carboxy-terminal autoprocessing domain. Cell 86, 21-34 (1996).

3. Lewis PM, et al. Cholesterol modification of sonic hedgehog is required for long-range signaling activity and effective modulation of signaling by Ptc1. Cell 105, 599-612 (2001).

4. Huang X, Litingtung Y, Chiang C. Region-specific requirement for cholesterol modification of sonic hedgehog in patterning the telencephalon and spinal cord. Development 134, 2095-2105 (2007).

5. Lee JD, et al. An acylatable residue of Hedgehog is differentially required in Drosophila and mouse limb development. Dev Biol 233, 122-136 (2001).

6. Corrales JD, Rocco GL, Blaess S, Guo Q, Joyner AL. Spatial pattern of sonic hedgehog signaling through Gli genes during cerebellum development. Development 131, 5581-5590 (2004).

7. Cordero D, Marcucio R, Hu D, Gaffield W, Tapadia M, Helms JA. Temporal perturbations in sonic hedgehog signaling elicit the spectrum of holoprosencephaly phenotypes. J Clin Invest 114, 485-494 (2004).

8. Dessaud E, et al. Interpretation of the sonic hedgehog morphogen gradient by a temporal adaptation mechanism. Nature 450, 717-720 (2007).

9. Garcia-Morales D, Navarro T, Iannini A, Pereira PS, Miguez DG, Casares F. Dynamic Hh signalling can generate temporal information during tissue patterning. Development 146, (2019).

10. Harfe BD, Scherz PJ, Nissim S, Tian H, McMahon AP, Tabin CJ. Evidence for an expansion-based temporal Shh gradient in specifying vertebrate digit identities. Cell 118, 517-528 (2004).

11. Nahmad M, Stathopoulos A. Dynamic interpretation of hedgehog signaling in the Drosophila wing disc. PLoS Biol 7, e1000202 (2009).

12. Ehring K, et al. Conserved cholesterol-related activities of Dispatched 1 drive Sonic hedgehog shedding from the cell membrane. J Cell Sci 135, (2022).

13. Coupland CE, et al. Structure, mechanism, and inhibition of Hedgehog acyltransferase. Mol Cell 81, 5025-5038 e5010 (2021).

14. Sacks FM, Jensen MK. From High-Density Lipoprotein Cholesterol to Measurements of Function: Prospects for the Development of Tests for High-Density Lipoprotein Functionality in Cardiovascular Disease. Arterioscler Thromb Vasc Biol 38, 487-499 (2018).

Changes from the earlier version to the revised version

Abstract

Summary

Earlier version: The Sonic hedgehog (Shh) signaling pathway controls embryonic development and tissue homeostasis after birth. This requires regulated solubilization of dual-lipidated, firmly plasma membrane-associated Shh precursors from producing cells. Although it is firmly established that the resistance-nodulation-division transporter Dispatched (Disp) drives this process, it is less clear how lipidated Shh solubilization from the plasma membrane is achieved. We previously showed that Disp enhances proteolytic Shh solubilization from its lipidated terminal peptide anchors. This process, called shedding, converts tightly membrane-associated hydrophobic Shh precursors into delipidated soluble proteins. We show here that Disp-mediated Shh shedding is modulated by a serum factor that we identify as high-density lipoprotein (HDL). In addition to serving as soluble sinks for free membrane cholesterol, HDLs also accept the cholesterol-modified Shh peptide from Disp. The cholesteroylated Shh peptide is required and sufficient for Disp-mediated transfer because mCherry linked to cholesteroylated peptides associates with HDL in a Disp-dependent manner, but an N-palmitoylated Shh variant that lacks C-cholesterol does not. Disp-mediated Shh transfer to HDL is finalized by proteolytic processing of the palmitoylated N-terminal membrane anchor. The resulting mono-lipidated Shh variant may help meet the demands for Hh activity regulation in different cell types and developing tissues.

Revised version: The Sonic hedgehog (Shh) signaling pathway controls embryonic development and tissue homeostasis after birth. This requires regulated solubilization of dual-lipidated, firmly plasma membrane-associated Shh precursors from producing cells. Although it is firmly established that the resistance-nodulation-division transporter Dispatched (Disp) drives this process, it is less clear how lipidated Shh solubilization from the plasma membrane is achieved. We previously showed that Disp enhances proteolytic Shh solubilization from its lipidated terminal peptide anchors. This process, called shedding, converts tightly membrane-associated hydrophobic Shh precursors into delipidated soluble proteins. We show here that Disp-mediated Shh shedding is modulated by a serum factor that we identify as high-density lipoprotein (HDL). In addition to serving as soluble sinks for free membrane cholesterol, HDLs also accept the cholesterolmodified Shh peptide from Disp. The cholesteroylated Shh peptide is required and sufficient for Disp-mediated transfer because mCherry linked to cholesteroylated peptides associates with HDL in a Disp-dependent manner, but an N-palmitoylated Shh variant that lacks C-cholesterol does not. Disp-mediated Shh transfer to HDL is finalized by proteolytic processing of the palmitoylated N-terminal membrane anchor. Unlike dual processed soluble Shh with moderate bioactivity, HDL-associated N-processed Shh is highly bioactive. We suggest that the purpose of different post-translationally modified soluble Shh forms generated from one dual-lipidated cellular precursor is to fine-tune cellular responses in a tissue type-specific and time-specific manner.

Introduction

Earlier version: Hedgehog (Hh) ligands activate an evolutionarily conserved signaling pathway that provides instructional cues during tissue morphogenesis, which, if misregulated, contributes to developmental disorders and cancer. Hhs are unique in that they require autocatalytic covalent cholesteroylation of the C-terminus 1 and N-terminal palmitoylation by a separate Hh acyltransferase (Hhat) activity 2. Both lipids firmly tether Hh to the plasma membrane of producing cells to effectively prevent unregulated ligand release 3. Signaling at distant target cells therefore requires the regulated solubilization of Hh from the membrane of ligand-producing cells, a process that is facilitated by vertebrate and invertebrate Dispatched (Disp) orthologs in vitro 4 and in vivo 5, 6, 7, 8. Disp contains 12 transmembrane helices and 2 extracellular domains and belongs to the resistance-nodulation-division family of transmembrane efflux pumps. Normally, such pumps maintain cellular homeostasis and remove toxic compounds. In addition, Disp contains a conserved multipass transmembrane domain known as the sterol-sensing domain (SSD) that regulates the homeostasis of free or esterified cellular cholesterol in other SSD proteins 9. This molecular architecture is in line with Disp extraction of free cholesterol from the plasma membrane to remove it from the cell 4. In addition, given its established role in Hh/Shh release, Disp may also extract the C-terminal Shh sterol to solubilize Hh/Shh into the extracellular compartment. Both possibilities are currently intensely investigated.

Revised version: Hedgehog (Hh) ligands activate an evolutionarily conserved signaling pathway that provides instructional cues during tissue morphogenesis, which, if misregulated, contributes to developmental disorders and cancer. Hhs are unique in that they require autocatalytic covalent cholesteroylation of the C-terminus 1 and N-terminal palmitoylation by a separate Hh acyltransferase (Hhat) activity 2. Both lipids firmly tether Hh to the plasma membrane of producing cells to effectively prevent unregulated ligand release 3. Signaling at distant target cells therefore requires the regulated solubilization of Hh from the membrane of ligand-producing cells, a process that is facilitated by vertebrate and invertebrate Dispatched (Disp) orthologs in vitro 4 and in vivo 5, 6, 7, 8. Disp contains 12 transmembrane helices and 2 extracellular domains and belongs to the resistance-nodulation-division family of transmembrane efflux pumps. Normally, such pumps maintain cellular homeostasis and remove toxic compounds. In addition, Disp contains a conserved multipass transmembrane domain known as the sterol-sensing domain (SSD) that regulates the homeostasis of free or esterified cellular cholesterol in other SSD proteins 9. This molecular architecture is in line with Disp extraction of free cholesterol from the plasma membrane to remove it from the cell 4. In addition, given its established role in Hh/Shh release, Disp has also been suggested to extract the C-terminal Shh sterol to solubilize Hh/Shh into the extracellular compartment 10.

Earlier version: Notably, Disp activity alone is insufficient to solubilize the vertebrate Hh family member Shh from the plasma membrane. A second required synergistic factor for maximum Shh signaling is the soluble extracellular glycoprotein Scube2 (Signal sequence, cubulin (CUB) domain, epidermal growth factor (EGF)-like protein 2) 9. One way to explain Disp/Scube2 synergy is by Disp-mediated extraction of dual-lipidated Shh from the plasma membrane to hand it over to Scube2, a mechanism that depends on the C-terminal cysteine-rich and CUB domains of Scube2 10, 11. In vitro support for this mechanism comes from Disp and Scube2 co-immunoprecipitation with Shh 10, 11 and from structural data 12. Other suggested carriers for Disp-extracted lipidated Hh/Shh include lipoproteins (LPPs) 13, 14, 15, exosomes as carriers of internalized and re-secreted Hhs 16, 17, 18, and micellous Hh assemblies 19. Another suggested Shh release mode is Disp-regulated proteolytic processing (called shedding) from the plasma membrane 4 by two major plasma-membrane-associated sheddases called A Disintegrin and Metalloproteinase (ADAM) 10 and 17 20, 21, 22. The role of Scube2 in this model is to enhance shedding of both terminal lipidated Shh peptides in a CUB domain-dependent manner 23, 24, 25. Consistent with this, CUB function in regulated shedding of other substrates often involves substrate recognition and induced structural changes in the substrate to boost turnover 25, 26, 27.

Revised version: However, Disp activity alone is insufficient to solubilize the vertebrate Hh family member Shh from the plasma membrane. A second required synergistic factor for maximum Shh signaling is the soluble extracellular glycoprotein Scube2 (Signal sequence, cubulin (CUB) domain, epidermal growth factor (EGF)-like protein 2) 9. One way to explain Disp/Scube2 synergy is by Disp-mediated extraction of dual-lipidated Shh from the plasma membrane to hand it over to Scube2, a mechanism that depends on the C-terminal cysteine-rich and CUB domains of Scube2 10, 11. In vitro support for this mechanism comes from Disp and Scube2 co-immunoprecipitation with Shh 10, 11 and from structural data 12. Other suggested carriers for Dispextracted lipidated Hh/Shh include lipoproteins (LPPs) 13, 14, 15, exosomes as carriers of internalized and re-secreted Hhs 16, 17, 18, and micellous Hh assemblies 19. Another suggested Shh release mode is Disp-regulated proteolytic processing (called shedding) from the plasma membrane 4 by two major plasma-membrane-associated sheddases called A Disintegrin and Metalloproteinase (ADAM) 10 and 17 20, 21, 22. The role of Scube2 in the shedding model is to enhance proteolytic processing of both terminal lipidated Shh peptides in a CUB domaindependent manner 23, 24, 25. Consistent with this, CUB function in regulated shedding of other substrates often involves substrate recognition and induced structural changes in the substrate to boost turnover 25, 26, 27.

Earlier version: In this study, we systematically characterized the biochemical parameters of Disp-regulated Shh solubilization with all these models and parameters in mind. To this end, we used a unique bicistronic Hhat/Shh co-expression system to ensure that only dual-lipidated Shh as a substrate for Disp was produced and analyzed 23. We further avoided any protein tags to ensure undisturbed Shh modification, secretion, and interaction with Disp and other potential pathway components. Finally, we used unbiased biochemical methods to analyze Shh solubilization, size and lipidation status under different experimental conditions. First, we confirmed by SDS-PAGE/immunoblotting and reverse-phase (RP)-HPLC that Disp and Scube2 synergistically enhance Shh shedding into the cell culture medium 23. Notably, we also found that repeated careful washing of Disp- and Shh-expressing cells to remove all serum traces abolished Shh shedding, and that high serum amounts promoted the solubilization of a previously uncharacterized fully bioactive Shh variant that lacks its palmitoylated N-terminus but retains its cholesteroylated C-terminus. Size exclusion chromatography (SEC) analyses demonstrated co-elution of this novel C-cholesteroylated Shh variant with serum apolipoprotein A1 (ApoA1), the major protein component of the high-density lipoprotein (HDL) fraction in serum. Consistent with this observation, purified HDL restored Shh solubilization from Shh- and Disp-expressing washed cells. We also found that the most C-terminal cholesteroylated Shh peptide is sufficient for Disp-mediated protein transfer to HDL. In contrast, palmitoylated N-terminal Shh membrane anchors are not transferred by Disp, but undergo shedding to complete Shh transfer to HDL. Our updated Hh shedding model therefore unifies previously disparate models of Disp-, sheddase-, and LPP-mediated Shh solubilization into one comprehensive system. This comprehensive system is fully in line with published in vivo Disp functions 5, 6, 7, 8, 28, in vivo Scube2 functions 29, 30, 31, in vitro roles of the Scube2 CUB domain 23, 25, the Disp structure 32, 33, and required N-terminal – but not C-terminal – shedding during Hh solubilization in vivo 34, 35, 36. Our model is also in line with the established importance of Hh C-cholesterol for Hh association into “large punctate” structures that are visible by light microscopy 37, C-cholesterol-dependent Hh solubilization 10, 38, and established in vivo roles of LPPs in Hh biofunction 13, 14, 15. We suggest that the purpose of Disp-regulated generation of different post-translationally modified soluble Shh forms, as shown in this study, is to fine-tune cellular responses in different Shh-receiving tissues at different developmental time points.

Revised version: In this study, we systematically characterized the biochemical parameters of Dispregulated Shh solubilization with all these published models and parameters in mind. To this end, we used a unique bicistronic Hhat/Shh co-expression system to ensure that only dual-lipidated Shh as a substrate for Disp was produced and analyzed 23. We further avoided any protein tags to ensure undisturbed Shh modification, secretion, and interaction with Disp and other potential pathway components. Finally, we used unbiased biochemical methods to analyze Shh solubilization, size, and lipidation status under different experimental conditions. First, we confirmed by SDS-PAGE/immunoblotting and reverse-phase (RP)-HPLC that Disp and Scube2 synergistically enhance Shh shedding into the cell culture medium 23. We also found that repeated careful washing of Disp- and Shh-expressing cells to remove all serum traces abolished Shh shedding, and that high serum amounts promoted the solubilization of a previously unknown, highly bioactive Shh variant that lacks its palmitoylated N-terminus but retains its cholesteroylated C-terminus. Size exclusion chromatography (SEC) analyses demonstrated co-elution of this novel Shh variant with serum apolipoprotein A1 (ApoA1), the major protein component of the high-density lipoprotein (HDL) fraction in serum. Consistent with this observation, purified HDL restored solubilization of the N-processed, bioactive Shh variant from washed cells. We also determined that the most C-terminal cholesteroylated Shh peptide is sufficient for Disp-mediated protein transfer to HDL. In contrast, palmitoylated N-terminal Shh membrane anchors are not transferred by Disp, but undergo shedding to complete Shh transfer to HDL. Our updated Hh shedding model therefore unifies previously disparate models of Disp-, sheddase-, and LPP-mediated Shh solubilization into one comprehensive system. This comprehensive system is fully in line with published reports of in vivo Disp functions 5, 6, 7, 8, 28, in vivo Scube2 functions 29, 30, 31, in vitro roles of the Scube2 CUB domain 23, 25, the Disp structure 32, 33, and required N-terminal – but not C-terminal – shedding during Hh solubilization in vivo 34, 35, 36. Our model is also in line with the established importance of Hh C-cholesterol for Hh association into “large punctate” structures that are visible by light microscopy and may represent Hh-LPP complexes 37, 38, 39, C-cholesterol-dependent Hh spread 10, 40, and previously established in vivo roles of LPPs in Hh biofunction 13, 14, 15. We suggest that the Disp-produced, HDL-linked mono-lipidated Shh variant presented in this work helps to meet the demands for Hh activity regulation in specific cell types and developing tissues in vivo.

Results

Synergistic Disp and Scube2 function increases shedding of dual-lipidated Shh from the plasma membrane

Earlier version: To analyze Shh solubilization from the plasma membrane, we produced dual-lipidated, tightly plasma membrane-associated morphogens in Bosc23 control cells that express endogenous Disp 23 (called CRISPR non-targeting control (nt ctrl) cells in this study) and in Bosc23 cells made deficient in Disp function by CRISPR/Cas9 (Using previously characterized Disp-/- cells 4; experimental outlines are described in Fig. S1). Shh biosynthesis in both cell lines starts with the covalent attachment of a cholesterol moiety to the C-terminus of the Shh signaling domain 39. This reaction is autocatalytic and strictly coupled to the generation of 19 kDa Shh signaling domains from a 45 kDa precursor. Subsequent N-palmitoylation of cholesteroylated Shh/Hh requires a separate enzymatic activity encoded by Hh acyltransferase (Hhat) 40. Because Bosc23 cells do not express endogenous Hhat 23, we minimized the production of non-palmitoylated Shh by using bicistronic mRNA that encodes the Shh precursor together with Hhat. Unlike Shh expression in the absence of Hhat, coupled Shh/Hhat expression ensures near-quantitative Shh N-palmitoylation in transfected Bosc23 cells 23. SDS-PAGE/immunoblotting was then used to characterize dual-lipidated Shh extraction from the plasma membrane of both cell lines into serum-depleted media (Fig. 1A; Fig. S2 includes all transfection controls and loading controls). We confirmed that Scube2 enhances Shh solubilization from nt ctrl cells (arrowhead) and that Disp-/- cells do not solubilize Shh 4 (Fig. 1A’). We also confirmed that the electrophoretic mobility of most Shh released from nt ctrl cells was increased over that of the corresponding dual-lipidated cellular material on the same blot 4, 23, 24 (asterisk). RP-HPLC of the solubilized material shown in Fig. 1A demonstrated that the observed electrophoretic mobility increase was caused by proteolytic removal of both lipidated terminal peptides (called shedding throughout this manuscript, Fig. 1A’’; see Fig. S2 for cellular dual lipidated or artificially produced monolipidated and unlipidated Shh standard proteins). Transgenic Disp expression in Disp-/- cells restored Shh shedding and solubilization, confirming the specificity of the assay (Fig. S2E) 4. Strikingly, the same phenotypic reversal was achieved by the co-expression of the cholesterol pump Patched (Ptc) 41 that depletes the plasma membrane of cholesterol 4. This suggests an indirect “second messenger” role of plasma membrane cholesterol not only in the regulation of Smoothened downstream of Ptc 41, 42, 43, but also in the regulation of Shh shedding downstream of Disp (Fig. S2E) 4. Note that alternative solubilized Shh variants (represented by the “upper” bands) were also present in the media. These variants, however, were always produced independent of Disp and Scube2 function. In contrast, Scube2 and Disp synergistically and specifically increased Shh shedding from the cell surface 4.

Revised version: To analyze Shh solubilization from the plasma membrane, we produced dual-lipidated, tightly plasma membrane-associated morphogens in Bosc23 control cells that express Disp endogenously 23 (called CRISPR non-targeting control (nt ctrl) cells in this study) and in Bosc23 cells made deficient in Disp function by CRISPR/Cas9 (using previously characterized Disp-/- cells 4; experimental outlines are described in Fig. S1). Shh biosynthesis in both cell lines starts with the covalent attachment of a cholesterol moiety to the C-terminus of the Shh signaling domain 41. This reaction is autocatalytic and strictly coupled to the generation of 19 kDa Shh signaling domains from a 45 kDa precursor. Subsequent N-palmitoylation of cholesteroylated Shh/Hh requires a separate enzymatic activity encoded by Hhat 42. Because Bosc23 cells do not express endogenous Hhat 23, we minimized the production of non-palmitoylated Shh by using bicistronic mRNA that encodes the Shh precursor together with Hhat. Unlike Shh expression in the absence of Hhat, coupled Shh/Hhat expression ensures near-quantitative Shh N-palmitoylation in transfected Bosc23 cells 23. SDS-PAGE/immunoblotting was then used to characterize dual-lipidated Shh extraction from the plasma membrane of both cell lines into serum-depleted media (Fig. 1A; Fig. S2 includes all transfection and loading controls). We confirmed the published finding that Scube2 enhances Shh solubilization from nt ctrl cells (arrowhead) and that Shh solubilization from Disp-/- cells was strongly reduced 4, 10, 11 (Fig. 1A’, Table S1). We also confirmed that the electrophoretic mobility of most Shh released from nt ctrl cells was increased over that of the corresponding dual-lipidated cellular material on the same blot 4, 23, 24 (asterisk). RP-HPLC of the solubilized material shown in Fig. 1A demonstrated that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both lipidated terminal peptides (this post-translational Shh modification is called shedding throughout this manuscript, Fig. 1A’’; see Fig. S2 for cellular dual-lipidated or artificially produced monolipidated and unlipidated Shh standard proteins and see reference 4, which ruled out alternative Shh deacylation modes). Transgenic Disp expression in Disp-/- cells restored Shh shedding and solubilization, confirming the specificity of the assay (Fig. S2E) 4. Strikingly, the same phenotypic reversal was achieved by the co-expression of the cholesterol pump Patched (Ptch1) 43, which depletes the plasma membrane of free sterols 4. This suggests an indirect “second messenger” role of plasma membrane cholesterol not only in the regulation of Smoothened downstream of Ptch1 43, 44, 45, but also in the regulation of Shh shedding by Disp (Fig. S2E) 4. Note that alternative solubilized Shh variants (represented by the “upper” bands) are also present in the media. Release of these variants, however, is independent of Disp and Scube2. In contrast, Scube2 and Disp synergistically and specifically increased Shh shedding from the cell surface (as indicated by the appearance of “lower” sized bands compared with those of cellular proteins) 4.

Disp and Scube2 enhance shedding of cell surface-associated dual-lipidated Shh into delipidated soluble forms.

Earlier version: Media containing 10% serum were exchanged for serum-free DMEM 36 h post-transfection and proteins solubilized for 6 h. Cells were not washed between media changes to leave residual serum traces in the assay. We refer to this experimental condition as “serum-depleted” throughout this work. A) Cap-dependent Shh translation and cap-independent Hhat translation from one bicistronic mRNA generated dual-lipidated, plasma membrane-associated proteins (asterisk in the cellular (c) fraction) in all transfected cells. Disp and Scube2 synergistically and specifically enhance dual-lipidated Shh precursor conversion into truncated soluble Shh (arrowhead). m: media, c: cell lysate. A’) Quantification of relative Shh release from nt ctrl and Disp-/- cells in the presence of Scube2. Ratios of solubilized versus cellular Shh were determined and expressed relative to Shh solubilization from nt ctrl cells (set to 100%, black bar). A’’) Reverse-phase HPLC analyses by using a C4 column revealed that specifically solubilized Shh (the same fraction indicated by the arrowhead in A, black line) was less hydrophobic than its cell surface-associated precursor (gray line, the asterisk denotes analysis of the same cellular fraction as shown in A). C4 column calibration and color-coding of fractions is explained in Fig. S2. B-D) Solubilization of non-palmitoylated C25SShh, non-cholesteroylated ShhN and lipid-free control C25SShhN under the same serum-depleted conditions. Arrowheads denote the Disp- and Scube2-solubilized material. B’-D’) Quantifications of B-D, again from nt ctrl and Disp-/- cells in the presence of Scube2. B’’-D’’) Reverse-phase HPLC demonstrates similar elution of C25SShh, ShhN, and non-lipidated C25SShhN. This suggests that lipidated peptide termini were removed during C25SShh and ShhN (as well as Shh, A’’) release. Unpaired t-test, two-tailed. ****: p<0.0001, *: p<0.016, n.t.: 0.79. n=20 datasets (A’), n=13 datasets (B’), n=10 datasets (C’), and n=17 datasets (D’). See Table S1 for detailed statistical information.

Revised version: Media containing 10% serum were exchanged for serum-free DMEM 36 h post-transfection and proteins solubilized for 6 h. Cells were not washed between media changes to leave residual serum traces in the assay. We refer to this experimental condition as “serum-depleted” throughout this work. A) Cap-dependent Shh translation and cap-independent Hhat translation from one bicistronic mRNA ensured the generation of dual-lipidated, plasma membrane-associated proteins (asterisk in the cellular (c) fraction) in all transfected cells. Disp and Scube2 synergistically and specifically enhance dual-lipidated Shh precursor conversion into truncated soluble variants during release (arrowhead). m: media, c: cell lysate. A’) Quantification of relative Shh release from nt ctrl and Disp-/- cells in the presence of Scube2. Ratios of solubilized versus cellular Shh were determined and expressed relative to Shh solubilization from nt ctrl cells (set to 100%, black bar). A’’) Reverse-phase HPLC analyses revealed that Shh solubilized by Disp and Scube2 (the same fraction indicated by the arrowhead in A, black line) was less hydrophobic than its cell surface-associated precursor (gray line; the asterisk denotes analysis of the same cellular fraction as shown in A). RP-HPLC calibration and color coding of Shh fractions is shown in Fig. S2 F-K. B-D) Solubilization of non-palmitoylated C25SShh (in this artificial variant, the N-terminal palmitate acceptor cysteine is changed for a non-accepting serine; functionally equivalent constructs with the cysteine exchanged for a non-accepting alanine (C25AShh) were also used in our study), non-cholesteroylated yet palmitoylated ShhN and lipid-free control C25SShhN under the same serum-depleted conditions. Arrowheads denote Shh variants that were solubilized by Disp- and Scube2. B’-D’) Quantifications of B-D, again from nt ctrl and Disp-/- cells in the presence of Scube2. B’’-D’’) Reverse-phase HPLC demonstrates similar elution of C25SShh, ShhN, and non-lipidated C25SShhN. This demonstrates that terminal lipids were removed during C25SShh and ShhN (as well as Shh, A’’) release. Unpaired t-test, two-tailed. ****: p<0.0001, *: p<0.016, n.t.: 0.79. n=20 datasets (A’), n=13 datasets (B’), n=10 datasets (C’), and n=17 datasets (D’). See Table S1 for detailed statistical information.

Earlier version: We also observed that solubilization of artificially produced, monolipidated Shh variants (C25SShh that lacks N-palmitate, Fig. 1B, B’, and ShhN that lacks C-cholesterol, Fig. 1C, C’, Fig. S2) was much less controlled, consistent with the observations of others 10, 11. Release of both monolipidated Shh variants remained linked to shedding of the respective lipidated membrane termini, as shown by electrophoretic mobility shifts of soluble proteins (Fig. 1B,C, arrowheads) and their delipidation during solubilization (Fig. 1B’’,C’’). Non-lipidated C25SShhN was always solubilized in a Disp- and Scube2-independent manner, as expected (Fig. 1D-D’’). These results suggest that cell-surface shedding represents a “ground-state” from which dual-lipidated Shh is protected, but only until Disp and Scube2 render it prone to shedding. Our results also explain the complete conservation of dual N- and C-terminal Hh lipidation during biosynthesis by the prevention of unregulated protein release from producing cells. This also demonstrates the importance of coupled Shh/Hhat expression to reliably characterize Disp- and Scube2-regulated Shh solubilization in vitro.

Revised version: We also observed that solubilization of artificially produced monolipidated Shh variants (C25SShh that lacks N-palmitate, Fig. 1B, B’, and ShhN that lacks C-cholesterol, Fig. 1C, C’, Fig. S2) was much less controlled, consistent with the observations of others 10, 11. Release of both monolipidated Shh variants remained linked to shedding of the respective lipidated membrane termini, as shown by electrophoretic mobility shifts of soluble proteins (Fig. 1B,C, arrowheads) and their delipidation during solubilization (as determined by RP-HPLC, Fig. 1B’’,C’’). Nonlipidated C25SShhN was always secreted in a Disp- and Scube2-independent manner, as expected (Fig. 1D-D’’). These results suggest that cell-surface shedding represents a “ground-state” from which dual-lipidated Shh is protected, but only until Disp and Scube2 render it prone to shedding. They also suggest that the complete conservation of dual N- and C-terminal Hh lipidation during biosynthesis serves to prevent unregulated protein release from producing cells. This demonstrates the importance of coupled Shh/Hhat expression in vitro to reliably characterize the mechanism of Disp- and Scube2-regulated Shh solubilization.

Earlier version: Shh solubilization depends on cleavage-activated Disp and the presence of serum

Revised version: Shh shedding depends on cleavage-activated Disp and the presence of serum

Earlier version: It was recently shown that the prohormone convertase furin cleaves Disp at a conserved processing site to activate the protein 44. On the basis of this activation mode, we hypothesized that furin inhibition should specifically impair Disp-regulated Shh shedding. To test this hypothesis, we added 50 μM peptidyl chloromethylketone (CMK, a competitive inhibitor of furin) to our solubilization assays. Indeed, CMK strongly reduced Shh shedding from the cell surface (Fig. 2A,B, Fig. S3A). During these assays, we also noted that repeated careful washing of cells to quantitively remove all serum traces strongly impaired Disp- and Scube2-mediated Shh solubilization into serum-free media (Fig. 2C, Fig. S3B). We derive two important conclusions from this latter observation. The first is that the minimal requirements of Na+-driven, Disp-mediated Shh extraction and hand-over to Scube2 45 are insufficient to explain Shh release. The second is that Shh self-assembly by the law of mass action 46 is also not supported, because this process should solubilize Shh regardless of the presence or absence of serum. In contrast, one function of Disp is to extract and transfer plasma membrane cholesterol to a soluble sink to transport it away from the cell 4, 47. In vertebrates, HDLs represent soluble sinks for “free” peripheral cholesterol 48, and HDLs as well as the pharmacological cholesterol chelator methyl--cyclodextrin (m cd) were previously shown to increase Shh shedding 4, 23, 49. These findings suggest that the permissive factor lacking from our shedding assay was most likely a soluble cholesterol acceptor, such as HDL or a related serum LPP.

Revised version: It was recently shown that the prohormone convertase furin cleaves Disp at a conserved processing site to activate it and to release Shh from the cell surface 46. On the basis of this activation mode, we hypothesized that furin inhibition may specifically impair Disp-regulated Shh shedding. To test this hypothesis, we added 0-80 μM peptidyl chloromethylketone (CMK, a competitive inhibitor of furin) to our solubilization assays. Indeed, CMK reduced Shh shedding from the cell surface in a concentration-dependent manner (Fig. 2A,B, Fig. S3A). During these assays, we also noted that repeated careful washing of cells to quantitively remove all serum traces strongly impaired Disp- and Scube2-mediated Shh solubilization into serum-free media (Fig. 2C, Fig. S3B). We derive two important conclusions from this latter observation. The first is that the minimal requirements of Na+-driven, Disp-mediated Shh extraction and handover to Scube2 47 are insufficient to release Shh. The second is that Shh self-assembly by the law of mass action 48 is also not supported, because this process should solubilize Shh regardless of the presence or absence of serum. In contrast, we previously found that one function of Disp is to extract and transfer free plasma membrane cholesterol to a soluble sink to transport it away from the cell 4, 49. In vertebrates, HDLs represent soluble sinks for “free” peripheral cholesterol 50, and HDLs and the pharmacological cholesterol chelator methyl-β-cyclodextrin have previously been shown to increase Shh shedding 4, 23, 51. Together, these findings suggested that the permissive factor lacking from our shedding assay was most likely a soluble cholesterol acceptor, such as HDL or a related serum LPP, as reported previously 14.

Shh shedding depends on Disp activation by furin and the presence of serum traces.

Earlier version: A) nt ctrl cells were transfected with Shh and Scube2, and Shh solubilization was monitored in the presence or absence of 50 μM CMK furin inhibitor. B) Quantification of CMK-inhibited Shh shedding. Ratios of solubilized versus cellular Shh were determined and expressed relative to Shh solubilization in the absence of inhibitor (set to 100%, black bar). Unpaired t-test, two-tailed. ****: p<0.0001, n=9. See Table S1 for additional statistical information. C) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum (referred to as “serum-free” conditions throughout this work). Note that Shh solubilization is completely abolished and that cellular precursors accumulate in the cells.

Revised version: A) nt ctrl cells were transfected with Shh and Scube2, and Shh solubilization was monitored in the presence or absence of CMK furin inhibitor. CMK impaired proteolytic processing and release of truncated (arrowheads) soluble Shh in a concentration-dependent manner. Right: Shh release in the presence or absence of 50 μM CMK furin inhibitor. B) Quantification of CMK-inhibited Shh shedding. Ratios of solubilized versus cellular Shh in the presence of 50 μM CMK inhibitor were determined and expressed relative to Shh solubilization in the absence of inhibitor (set to 100%, black bar). Unpaired t-test, twotailed. ****: p<0.0001, n=9. See Fig. S3A for loading controls and Table S1 for additional statistical information. C) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum (referred to as “serum-free” conditions throughout this work). Note that Shh solubilization under serum-free conditions is abolished.

Increased serum concentrations shift shedding of both Shh termini to selective N-terminal shedding

Earlier version: To characterize serum-enhanced shedding in more detail, dual-lipidated Shh was expressed in nt ctrl cells and Disp-/- cells and solubilized into serum-depleted Dulbecco’s Modified Eagle’s Medium (DMEM) or into DMEM supplemented with 0.05%, 2%, and 10% fetal calf serum (FCS) (Figs. 3A-D; Fig. S3 shows transfection and loading controls). We observed that Scube2 enhances Shh shedding from Disp-expressing cells into serum-depleted media (Fig. 3A, arrowhead), but does not enhance Shh shedding from Disp-/- cells. RP-HPLC of the solubilized material confirmed that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both terminal lipidated peptides (Fig. 3A’). Increased serum concentrations during Shh solubilization (0.05% FCS (Fig. 3B), 5% FCS (Fig. 3C), and 10% FCS (Fig. 3D)) did not appear to affect Disp- and Scube2-specific Shh shedding much (arrowheads). To our surprise, however, RP-HPLC of the solubilized materials shown in Fig. 3B-D revealed a gradual shift from dual Shh shedding (labeled [1] in Fig. 3A’-D’) toward the solubilization of variant proteins with their C-terminal cholesteroylated peptide still intact (labeled [2] in Fig. 3B’-D’, Fig. S4A-E). We note that dual-lipidated Shh was also present in serum-containing media (indicated by an asterisk in Fig. S4A-E). However, we found that relative amounts of this fraction increased strongly in the absence of Scube2 (Fig. S4F; here, cells required incubation for 24 h to compensate for low Shh solubilization) and in the absence of Disp-/- (Fig. S4G), while Scube2 expression in Disp-/- cells increased relative amounts of delipidated soluble Shh (Fig. S4H). Disp- and Scube2-independent solubilization of dual-lipidated Shh, as observed here (Fig. 1A), indicates that it is probably an overexpression artefact or caused by soluble plasma membrane remnants of apototic cells.

Revised version: To characterize serum-dependent shedding in more detail, we expressed dual-lipidated Shh in nt ctrl cells and Disp-/- cells and solubilized the proteins into serum-depleted Dulbecco’s Modified Eagle’s Medium (DMEM) or into DMEM supplemented with 0.05%, 5%, and 10% fetal calf serum (FCS) (Fig. 3A-D; Fig. S3C-F shows transfection and loading controls). Consistent with previous observations, we again found that Scube2 enhances Shh shedding from Disp-expressing cells into serum-depleted media (Fig. 3A, arrowhead), but does not much enhance Shh shedding from Disp-/- cells. RP-HPLC of the solubilized material confirmed that the observed increase in electrophoretic mobility was caused by the proteolytic removal of both terminal lipidated peptides (Fig. 3A’). Increased serum concentrations during Shh solubilization (0.05% FCS (Fig. 3B), 5% FCS (Fig. 3C), and 10% FCS (Fig. 3D)) did not appear to much affect Disp- and Scube2-specific Shh shedding (arrowheads). To our surprise, however, RP-HPLC of the solubilized materials shown in Fig. 3B-D revealed a gradual shift from dual Shh shedding (labeled [1] in Fig. 3A’-D’) toward the solubilization of variant proteins with their C-terminal cholesteroylated peptide still intact (labeled [2] in Fig. 3B’-D’, Fig. S4A-E). We note that low amounts of dual-lipidated Shh were also present in serum-containing media (indicated by asterisks in Fig. S4B-E). However, relative amounts of this protein fraction increased both in the absence of Scube2 (Fig. S4F; here, cells required incubation for 24 h to compensate for low Shh solubilization) and in the absence of Disp (Fig. S4G). In contrast, Scube2 expression in Disp-/- cells increased relative amounts of delipidated soluble Shh (Fig. S4H). We conclude from the Disp- and Scube2-independent solubilization of dual-lipidated Shh that it is physiologically irrelevant.

Dual-lipidated cell-surface Shh converts into delipidated soluble forms under low-serum and high-serum conditions.

Earlier version: Media were changed for serum-free DMEM 36 h post-transfection (cells were not washed) or DMEM containing the indicated amounts of serum, and proteins were solubilized for 6 h (serum-depleted) or 24h (with 0.05%, 5% and 10% serum). A) Under serum-depleted conditions, Disp and Scube2 increase dual-lipidated Shh shedding into truncated soluble forms (arrowhead). m: media, c: cell lysate. A’) RP-HPLC confirmed loss of both terminal lipidated Shh peptides during Disp- and Scube2-regulated shedding. B-D) The appearance of truncated Shh in serum-containing media remained dependent on Disp and Scube2 (arrowheads). B’-D’) RP-HPLC revealed that increased serum amounts shift dual Shh shedding (low serum amounts, [1]) toward N-terminally restricted shedding and release of a cholesteroylated Shh form [2]. Low levels of dual-lipidated Shh are also detected [3]. See Fig. S4 for experimental details.

Revised version: Media were changed for serum-free DMEM 36 h post-transfection (cells were not washed) or DMEM containing the indicated amounts of serum, and proteins were solubilized for 6 h (serum-depleted) or for 24 h (with 0.05%, 5%, and 10% serum). A) Under serum-depleted conditions, Disp and Scube2 increase dual-lipidated Shh shedding into truncated soluble forms (arrowhead). m: media, c: cell lysate. A’) RP-HPLC confirmed loss of both terminal lipidated Shh peptides during Disp- and Scube2-regulated shedding. B-D) The appearance of truncated Shh in serum-containing media remained dependent on Disp and – to a lesser degree – on Scube2 (arrowheads). B’-D’) RP-HPLC revealed that increased serum amounts shift dual Shh shedding (low serum amounts, [1]) toward N-terminally restricted shedding and release of a cholesteroylated Shh form [2]. Low levels of dual-lipidated Shh are also detected [3]. See Fig. S4 for experimental details.

Earlier version: So far, our data supported that Scube2 increases N-terminal and C-terminal Shh shedding from Disp-expressing cells and that serum enhances this process, likely by providing a sink for membrane cholesterol transferred by Disp 4 (Fig. S4I). Moreover, high serum amounts seem to promote a second Disp solubilization mode in which Shh shedding is restricted to the palmitoylated N-peptide while the C-terminus remains intact. This latter finding raised the interesting possibility that serum factors may accept or protect the cholesteroylated Shh C-terminus. Note that a filtrate of blood serum through the capillary walls, called interstital fluid, represents the microenvironment of tissues and cells in vivo, also of those that express and solubilize Shh during development 14. Interstitial fluid is rich in proteins and lipoprotein particles of small mass from the serum 50. Therefore, we suggest that Shh expression in serum presence resembles conditions in Shh expressing tissues in vivo, and may therefore be physiologically relevant.

Revised version: So far, our data showed that Scube2 increases N-terminal and C-terminal Shh shedding from Disp-expressing cells and that serum enhances this process, likely by providing a sink for membrane cholesterol transferred by Disp 4 (Fig. S4I). We also showed that high serum amounts promote a second Disp solubilization mode in which Shh shedding is restricted to the palmitoylated N-peptide while the C-terminus remains intact. This latter finding raised the interesting possibility that serum factors may accept or protect the cholesteroylated Shh C-terminus. Indeed, it is known that a filtrate of blood serum through the capillary walls, called interstitial fluid, represents the microenvironment of tissues and cells in vivo, as well as of those that express and solubilize Shh during development 14. It is also known that interstitial fluid is rich in proteins and LPP particles of small mass from the serum 52. This suggests that Shh expression in the presence of serum may resemble the conditions in Shh-expressing tissues in vivo, and may therefore be closer to physiological relevance.

N-palmitate is not required for Shh-induced in vitro differentiation of C3H10T1/2 cells

Earlier version: Is this novel N-terminally shed Shh variant in serum-containing media functional? It is known that dual lipidation is required for unperturbed Hh biofunction in vivo 51, 52, 53, 54, 55, and dual-lipidated Shh expression in vitro generates soluble variants that are 10-30 times more bioactive than C25SShh that does not undergo Hhat-catalyzed N-palmitoylation during its biosynthesis, yet undergoes unperturbed C-terminal cholesteroylation and secretion to the cell surface (Fig. 1B) 40, 55. According to these published observations, the selectively N-terminally shed Shh variant described here should not be very active. To test this possibility, we used the multipotent fibroblastic C3H10T1/2 cell line as a reporter. We first verified multipotency of our C3H10T1/2 reporter cell line to differentiate into osteoblasts 56, chondrocytes 57 or adipocytes 58 (Fig. S5A). To this end, C3H10T1/2 cells were cultured for various time periods in the presence of adipogenic, chondrogenic and osteogenic supplements, and their responsiveness confirmed based on phenotype and the expression of cell surface markers. We then incubated C3H10T1/2 cells with R&D 8908-SH (Fig. 4A). R&D 8908-SH is commercially available, dual-lipidated Shh obtained by detergent-extraction from transfected cells and shares 91% sequence identity and both lipids with Indian hedgehog, an established osteogenic factor 59. R&D 8908-SH induced alkaline phosphatase (Alp) expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 56. We also incubated C3H10T1/2 cells with Shh expressed in the presence of serum and physiological release regulators Scube2 30, 31 and Disp 5, 6 (Fig. 4A, inset). Surprisingly, we found that the activity of solubilized Shh was increased over similar amounts of R&D 8908-SH (Fig. 4A, inset), and C25AShh expressed under the same conditions was equally active. Quantitative reverse transcription-polymerase chain reaction analysis (qPCR) of Ptc1, Ptc2 and Gli1 expression in Shh-treated C3H10T1/2 cells (Fig. 4B, Fig. S5B) confirmed this finding: After protein normalization (using 4ng R&D 8908-SH and 100μl Shh or C25SShh), qPCR confirmed similar mRNA expression increases of Ptc1 and Gli1 under direct Shh, C25SShh or R&D 8908-SH control (Fig. 4B). Ptc is known to be upregulated by Shh 60, and Gli1 is a zinc finger transcription factor that operates downstream of Ptc (mainly Ptc1 in vertebrates, a less characterized Ptc2 homolog is also expressed) and is also transcribed in a Shh-dependent manner 61. Apparent differences between determined Ptc2 levels, as observed in this assay, can be explained by the very low Ptc2 expression levels in C3H10T1/2 cells that give rise to high margins of error (Ptc2 Cq∼33 for Shh-treated cells, Ptc2 Cq∼31 for C25AShh-treated cells, Ptc2 Cq∼37 and ∼34 for R&D 8908-SH-treated cells). Key regulators for adipogenesis, osteogenesis, chondrogenesis and proliferation in Shh and C25SShh-induced C3H10T1/2 cells were also similar (Fig. S5B,C). Taken together, our results demonstrate that N-palmitate is not essential for Shh signaling if proteins were produced and released in the presence of serum.

Revised version: Is this novel N-terminally processed Shh variant fully functional? It is well-established that dual lipidation during biosynthesis is absolutely required for unperturbed Hh/Shh biofunction in vivo 39, 53, 54, 55, 56. These studies showed that dual-lipidated Hh/Shh expression in vitro generates soluble variants that are 10-30 times more bioactive than artificial proteins that do not undergo Hhat-catalyzed N-palmitoylation, yet undergo unperturbed C-terminal cholesteroylation and secretion to the cell surface (Fig. 1B) 42, 56. According to these published observations, the N-terminally processed Shh variant described here should not be very active because it also lacks the palmitate. To test this hypothesis, we used the Ptch1-expressing, multipotent fibroblastic C3H10T1/2 cell line as a reporter 57. We first verified multipotency of our C3H10T1/2 reporter cell line to differentiate into osteoblasts 57, chondrocytes 58, or adipocytes 59 (Fig. S5A). To this end, C3H10T1/2 cells were cultured for various time periods in the presence of adipogenic, chondrogenic, and osteogenic supplements, and their responsiveness was confirmed on the basis of phenotype and the expression of cell surface markers. We then incubated C3H10T1/2 cells with Shh expressed in the presence of 10% serum and physiological release regulators Scube2 30, 31 and Disp 5, 6 (Fig. 4A). Shh shares 91% sequence identity and both lipids with Indian hedgehog, an established osteogenic factor 60, and both Indian hedgehog and Shh stimulate C3H10T1/2 osteogenic differentiation. We confirmed that Shh induced alkaline phosphatase (Alp) expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57 (Fig. 4A). Notably, we observed that C25AShh expressed under the same conditions was equally active or even more active than Shh (Fig. 4A). Quantitative reverse transcription-polymerase chain reaction analysis (qPCR) of Ptch1 and Gli1 expression in Shh-treated C3H10T1/2 cells (Fig. 4B, Fig. S5B) or NIH3T3 cells (Fig. S5B) confirmed this finding: After protein normalization, qPCR confirmed similar increases in mRNA expression of Ptch1 and Gli1 under direct Shh or C25AShh control. Ptch1 is known to be upregulated by Shh 61, and Gli1 is a zinc finger transcription factor that operates downstream of Ptch1 and is also transcribed in a Hh-dependent manner 62. Key regulators for adipogenesis, osteogenesis, chondrogenesis, and proliferation in Shh and C25SShh-induced C3H10T1/2 cells were also similar (Fig. S5C,D). In contrast, C25SShhN negative control biofunction was always strongly impaired. Finally, we incubated C3H10T1/2 cells with high-activity Shh R&D 8908-SH, a commercially available, dual-lipidated Shh variant obtained by detergent extraction from transfected cells (Fig. S5E). We observed that R&D 8908-SH induced Alp expression in differentiating C3H10T1/2 cells in a concentration-dependent manner, as expected 57. We also found that the activities of solubilized Shh and C25AShh were increased over similar amounts of R&D 8908-SH. qPCR of Ptch1 and Gli1 expression confirmed high activities of R&D 8908-SH, Shh and C25AShh (Fig. S5F). These results demonstrate that N-palmitate is not essential for Shh signaling strength at the level of Ptch1 if Shh is released in the presence of serum.

Earlier version: Activities and SEC of dual-lipidated Shh/Hh and depalmitoylated Shh variants solubilized into serum-containing media. A) Immunoblotted proteins (inset) were used for quantification and visualization of dually lipidated R&D 8908-SH, Shh, C25AShh (this artificial variant has the cysteine palmitate acceptor changed for a non-accepting alanine) and unlipidated C25SShhN. Data were normalized to mock-treated C3H10T1/2 cells. Shh, C25AShh and C25SShhN were solubilized from Scube2- and Disp-expressing cells into media containing 10% FCS. 0-4 ng of R&D 8908-SH induced C3H10T1/2 reporter differentiation in a concentration-dependent manner, as determined by Alp activity measured at 405nm (blue line). Similar amounts of solubilized Shh, as determined by immunoblotting, increased C3H10T1/2 osteogenesis over that of the dual-lipidated R&D 8908-SH standard (black line). Notably, C25AShh was also highly bioactive if released into serum-containing medium and tested in this experimental system (green line). Inset: Immunoblot of proteins used in this assay. The arrow indicates proteolytic processing of the palmitoylated N-terminal peptide from Shh and C25AShh, but not from R&D 8908-SH. Artificial unlipidated soluble C25SShhN showed only little activity 63 (gray line). B) Similar induction of transcription of Hh target genes Ptc1 and Gli1 in C3H10T1/2 reporter cells. Cells were stimulated with 4ng R&D 8908-SH or similar amounts of Shh and C25AShh. Shh and C25AShh were both solubilized from Bosc23 cells into serum-containing media in the presence of Scube2 (100μl Shh/C25AShh+Scube2). C) SEC reveals substantial amounts of increased molecular weight Shh in media that contain 10% serum (black line). The increased molecular weight Shh eluted together with ApoA1 (orange line). In contrast, Scube2 was largely monomeric in solution (blue line). Levels of Shh-induced alkaline phosphatase activity in C3H10T1/2 cells were again measured as absorption at 405 nm, demonstrating the strongest C3H10T1/2 differentiation by eluted fractions containing large Shh assemblies. D) SEC of Drosophila Hh (black line) and of a variant lacking its HS binding site (HhHS, black dotted line). Both proteins were expressed under actin-Gal4/UAS-control from S2 cells and solubilized into media containing 10% FCS.

Revised version: Activities and SEC of dual-lipidated Shh/Hh and depalmitoylated Shh variants solubilized into serum-containing media. A) Shh, C25AShh (this artificial variant has the cysteine palmitate acceptor changed for a non-accepting alanine), and non-lipidated C25AShhN were expressed into media containing 10% FCS; their protein amounts determined by immunoblotting and normalized; and conditioned media added to C3H10T1/2 reporter cells to induce their Hh-dependent differentiation into alkaline phosphatase (Alp)-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At lower concentrations (1x), Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. At higher concentrations (2x), the bioactivity of C25AShh was increased over that of Shh. C25AShhN was inactive. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, n.s.=0.99, n=3-9. See Table S1 for additional statistical information and Fig. S5E,F, which demonstrates protein activities similar to, or exceeding, that of a dual-lipidated Shh control. B) Similar transcription of Hh target genes Ptch1 and Gli1 by Shh and C25AShh in two independent experiments. C3H10T1/2 reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and lower (1x) concentrations. C) SEC reveals substantial amounts of increased molecular weight Shh in media that contain 10% serum (black line). The increased molecular weight Shh eluted together with ApoA1 (orange line). In contrast, Scube2 was largely monomeric in solution (blue line). Levels of Shh-induced Alp activity in C3H10T1/2 cells were measured as absorption at 405 nm, demonstrating the strongest C3H10T1/2 differentiation by eluted fractions containing large Shh assemblies. D) SEC of Drosophila Hh (black line) and of a variant lacking its HS binding site (HhβHS, black dotted line). Both proteins were expressed under actin-Gal4/UAS-control from S2 cells and solubilized into media containing 10% FCS.

Soluble Shh/Hh and ApoA1-containing LPPs have similar sizes

Earlier version: A decade ago, it had already been shown that sterol-modified Hh/Shh is released in LPP-associated bioactive form – as wellas in desteroylated, non-associated form – in flies and mammals 14. This suggested that serum LPPs in our assays may have promoted dual Shh shedding 4 (Fig. 1) but possibly also promoted C-cholesteroylated Shh assembly into soluble complexes. We used size exclusion chromatography (SEC) to test this hypothesis. SEC detects soluble Shh monomers (20 kDa) and covers the entire molecular weight range up to a MW cutoff of about 1 mDa. The molecular weight span of serum HDL ranges from 1.75ξ105 Da to 3.6ξ105 Da, corresponding to particle sizes of 5-11 nm. This small size range makes HDLs abundant constituents of interstitial fluid and good candidates for Shh solubilization from cells and tissues. Larger complexes in mammalian serum include low-density lipoproteins (LDLs, MW 2.75ξ106 Da) and very low-density lipoproteins (VLDLs, MW 10-80ξ106 Da) that both elute in the void column volume. As shown in Fig. 4C, Shh in DMEM + 10% FCS elutes in fractions that cover the entire molecular weight range (black line). One prominent Shh peak is detected around 20 kDa – most likely representing the fully delipidated Shh fraction (Fig. 3, labeled [1]) – and a second peak is detected between 300 kDa and 600 kDa. Reprobing the same (stripped) blot with antibodies directed against ApoA1 (the major protein component of HDL) revealed an elution profile that overlapped with that of Shh (Fig. 4C, orange line). Of note, most Shh co-eluted with the largest HDL particles, suggesting that maturation of small early (discoidal) HDL into larger (spherical) HDL upon incorporation of cholesterol may be accompanied by gradual Shh transfer to the same particles. In contrast, Scube2 (Fig. 4C, blue line) co-eluted with smaller Shh multimers 46. In order to determine which size fraction contained the most biologically active Shh, we analyzed C3H10T1/2 osteoblast precursor cell differentiation by using aliquots of eluted fractions from the same SEC run. Shh-induced C3H10T1/2 cells in a size-dependent activity distribution: Large Shh assemblies were strongly bioactive, smaller assemblies were also strongly bioactive, but monomeric Shh was only moderately active (Fig. 4C, green dashed line). We also found that Drosophila Hh expressed in S2 cells and solubilized into serum-containing medium showed a similar size distribution (Fig. 4D, black line) and that Hh assemblies co-eluted with serum ApoA1 as well (orange line). This suggests that the mechanism of serum-enhanced Hh assembly during solubilization is conserved. The possibility of Hh assembly as a consequence of heparan sulfate (HS) proteoglycan interactions 62 was excluded because site-directed mutagenesis of the HS binding Cardin-Weintraub motif did not change HhHS size distribution in our assays (Fig. 4D, black dotted line). From these observations, we conclude that large Shh/Hh assemblies form by HDL association in vitro, which is in line with a previous study showing that Drosophila Hh copurifies with LPP and co-localizes with LPP in the developing wing epithelium in vivo 15. The same study showed that reduced LPP levels in larvae impair Hh release and lead to Hh accumulation at the site of production, similar to the in vitro findings shown in Fig. 2C.

Revised version: A decade ago, it was already known that flies and mammals release sterol-modified Hh/Shh in LPP-associated bioactive form, yet also in desteroylated, non-associated form 14. These findings suggested that serum LPPs in our assays not only may have promoted dual Shh shedding 4 (Fig. 1), but may also have promoted monolipidated Shh assembly into soluble complexes 14. We used size exclusion chromatography (SEC) to test this possibility. SEC detects soluble Shh monomers (20 kDa) and covers the entire molecular weight (MW) range up to a cutoff of about 1 mDa. The MW span of serum HDL ranges from 1.75×105 Da to 3.6×105 Da, corresponding to particle sizes of 5-11 nm. This small size range makes HDLs abundant constituents of interstitial fluid and good candidates for Shh solubilization from cells and tissues. Larger complexes in mammalian serum include low-density lipoproteins (LDLs, MW 2.75×106 Da) and very low-density lipoproteins (VLDLs, MW 10-80×106 Da) that would both elute in the void column volume. As shown in Fig. 4C, Shh in DMEM + 10% FCS elutes in fractions that cover the entire MW range (black line). One prominent Shh peak is detected around 20 kDa – most likely representing the fully delipidated dually shed Shh fraction (Fig. 3, labeled [1]) – and a second peak is detected between 300 kDa and 600 kDa. Reprobing the same (stripped) blot with antibodies directed against ApoA1 (the major protein component of HDL) revealed an elution profile that overlapped with that of Shh (Fig. 4C, orange line). The leftward shift of Shh elution relative to HDL elution may be explained by the increased size of a LPP subfraction after its Disp-mediated loading with Shh. In contrast, Scube2 (Fig. 4C, blue line) co-eluted only with smaller Shh multimers 48. In order to determine which size fraction contained biologically active Shh, we analyzed C3H10T1/2 osteoblast precursor cell differentiation by using aliquots of eluted fractions from the same SEC run. We found that Shh induced C3H10T1/2 differentiation into osteoblasts in a size-dependent activity distribution: Large Shh assemblies were strongly bioactive, and smaller assemblies were also bioactive, but monomeric Shh was only moderately active (Fig. 4C, green dashed line). We also found that Drosophila Hh expressed in S2 cells and solubilized into serum-containing medium showed a similar size distribution (Fig. 4D, black line) and that Hh assemblies also co-eluted with serum ApoA1 (orange line). This suggests that the mechanism of serum-enhanced Hh assembly during solubilization is conserved. The alternative possibility of Hh assembly as a consequence of heparan sulfate (HS) proteoglycan interactions 63 was dismissed because site-directed mutagenesis of the HS binding Cardin-Weintraub motif did not change HhΔHS size distribution (Fig. 4D, black dotted line). These findings confirm a previous study showing that Drosophila Hh copurifies with LPP and colocalizes with LPP in the developing wing epithelium in vivo 15. The same study showed that reduced LPP levels in larvae impair Hh release and lead to Hh accumulation at the site of production, similar to the in vitro findings shown in Fig. 2C.

Earlier version: Disp function depends on HDL

Revised version: Disp function requires HDL

Earlier version: To test whether soluble HDL is sufficient to increase Disp-mediated Shh solubilization, we next induced this process directly. To this end, we expressed Shh in nt ctrl cells and Disp-/- cells in the presence or absence of Scube2, washed the cells 36 h post transfection, and added serum-free DMEM or DMEM supplemented with 40 μg/mL purified human HDL 4. We first confirmed that lack of serum factors strongly reduced Shh solubilization (Fig. 5A, labeled [1], Fig. S6A). RP-HPLC of the small amounts of released Shh confirmed removal of both lipidated Shh peptide anchors during solubilization (Fig. 5A’, Fig. S6A’), and SEC confirmed that proteins were monomeric (Fig. 5A’’). In contrast, Shh solubilization from nt ctrl cells into DMEM supplemented with 40 μg/mL HDL was strongly increased (Fig. 5B,C, Fig. S6B). As previously observed for Shh expression in the presence of serum, RP-HPLC of HDL-solubilized Shh [2] revealed that its C-terminal cholesteroylated peptides were still intact (Fig. 5B’, Fig. S6B’), and that the proteins were strongly bioactive (Fig. 5B’’). Solubilized dual-lipidated Shh was also found, but its solubilization was independent of Scube2 and Disp function (Fig. S6B’’,B’’’), indicating unspecific release as observed before (Fig. 1). From the observed size shift and overlap with the ApoA1 elution profile in SEC – note that the observed Shh molecular weight range almost exactly matches that of HDL – we suggest that monolipidated truncated Shh associates with HDL in a Disp-regulated manner. This possibility is supported by the previous observation that Hh levels in the hemolymph of Disp-deficient fly larvae were strongly decreased 14. We also tested whether HDL constitutes the only possible Shh acceptor, or whether LDL can also carry monolipidated morphogens. As shown in Fig. S6C,D, Shh transferred to HDL-supplemented serum-free DMEM, but not to LDL-supplemented DMEM.

Revised version: To test whether the soluble LPP that increases Disp-mediated solubilization of N-processed, cholesteroylated Shh is HDL, we tested this possibility directly. To this end, we expressed Shh in nt ctrl cells and Disp-/- cells in the presence or absence of Scube2, washed the cells 36 h post-transfection to inhibit Shh release, and added serum-free DMEM or DMEM supplemented with 40 μg/mL purified human HDL 4. We first confirmed that lack of serum factors strongly reduced Shh solubilization (Fig. 5A, labeled [1], Fig. S6A). RP-HPLC of the small amounts of released Shh confirmed removal of both lipidated Shh peptide anchors during solubilization (Fig. 5A’, Fig. S6A’), and SEC confirmed that solubilized proteins were monomeric (Fig. 5A’’). In contrast, Shh solubilization from nt ctrl cells into DMEM supplemented with 40 μg/mL HDL was strongly increased (Fig. 5B,C, Fig. S6B). As previously observed for Shh expression in the presence of serum, RP-HPLC of HDL-solubilized Shh (Fig. 5B, labeled [2]) confirmed that its C-terminal cholesteroylated peptides were still intact (Fig. 5B’, Fig. S6B’) and proteins were strongly bioactive (Fig. 5B’’, green dashed line). Solubilized dual-lipidated Shh was also found, but its solubilization was again independent of Scube2 and Disp function (Fig. S6B’’,B’’’), indicating unspecific release. From the observed size shift and overlap with ApoA1 and ApoE4 elution profiles in SEC (note that the observed Shh MW range almost matches that of ApoA1), we suggest that monolipidated truncated Shh associates with HDL in a Disp-regulated manner. This possibility is supported by the previous observation that Hh levels in the hemolymph of Disp-deficient fly larvae were strongly decreased 14. We also tested whether HDL constitutes the only possible Shh acceptor, or whether LDL can also carry monolipidated morphogens. As shown in Fig. S6C,D, Shh transferred to HDL-supplemented serum-free DMEM, but not to LDL-supplemented serum-free DMEM. Shh released by HDL is N-terminally processed (Fig. S6G) and physically interacts with the LPP (Fig. S6E,F).

HDL increases N-processed Shh solubilization by Disp.

Earlier version: Media were changed for serum-free DMEM 36 h post-transfection (and cells washed three times) or serum-free DMEM supplemented with 40 mg/mL HDL before proteins were solubilized for 6 h. A) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum. Note that Shh solubilization is strongly reduced [1]. A’) RP-HPLC of the material labeled [1] revealed complete delipidation during release. A’’) SEC of the same delipidated material demonstrates that it is readily soluble and not associated with Scube2. B) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the presence of 40 μg/mL HDL. Shh shedding and solubilization are strongly increased by Disp [2] but not by Scube2. B’) RP-HPLC of the material labeled [2] revealed that HDL shifts Shh shedding from dual processing (A’, [1]) toward the release of cholesteroylated Shh. B’’) SEC of the same material [2] (black line) shows a molecular weight increase that matches the molecular weight range of HDL (orange line). Again, the soluble Shh elution profile does not overlap with that of Scube2 (blue line). C) Quantification of HDL-induced Shh solubilization from nt ctrl cells and Disp-/- cells. Unpaired t-test, two-tailed. ****: p<0.0001, n=6. See Table S1 for additional statistical information.

Revised version: Media were changed for serum-free DMEM 36 h post-transfection (and cells washed three times) or serum-free DMEM supplemented with 40 μg/mL HDL before proteins were solubilized for 6 h. A) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the complete absence of serum. Note that Shh solubilization is strongly reduced [1]. A’) RP-HPLC of the material labeled [1] revealed complete delipidation during release. A’’) SEC of the same delipidated material demonstrates that it is readily soluble and not associated with Scube2. B) Immunoblotted cellular (c) and medium (m) fractions of Shh expressing nt ctrl and Disp-/- cells in the presence of 40 μg/mL HDL. Shh shedding and solubilization are strongly increased by Disp [2] but not by Scube2. B’) RP-HPLC of the material labeled [2] revealed that HDL shifts Shh shedding from dual processing (A’, [1]) toward the release of cholesteroylated Shh. B’’) SEC of the same material [2] (black line) shows a molecular weight increase that matches the molecular weight range of HDL, as indicated by marker apolipoproteins ApoA1 (orange line) and mobile ApoE4. The former lends structural stability to the particle and stimulates cholesterol efflux to HDL; the latter facilitates cholesterol storage and core expansion and is therefore a marker of larger mature HDL particles (brown dashed line). Again, the soluble Shh elution profile does not overlap with that of Scube2 (blue line). C) Quantification of HDL-induced Shh solubilization from nt ctrl cells and Disp-/- cells. Unpaired t-test, two-tailed. ****: p<0.0001, n=6. See Table S1 for additional statistical information.

Added in the revised version: Is N-terminally processed, HDL-associated Shh functional? To answer this question, we again used the multipotent fibroblastic C3H10T1/2 cell line as a reporter. Shh was co-expressed with Scube2 in the presence of 80 μg/mL purified human HDL and the conditioned media added to C3H10T1/2 cells. As expected, Shh/HDL complexes induced Alp expression and C3H10T1/2 differentiation into osteoblasts in a concentration-dependent manner 57, and C25AShh expressed under the same conditions was equally active (Fig. 6A). qPCR of Ptch1 and Gli1 expression in Shh/HDL-stimulated or C25AShh/HDL-stimulated C3H10T1/2 cells confirmed this finding: Both protein/HDL complexes induced similar increases in mRNA expression of Ptch1 and Gli1 (Fig. 6B). This suggests full bioactivity of the HDL-associated, N-terminally truncated Shh variants that we describe in this study.

Added in the revised version: Activities of HDL-associated Shh and non-palmitoylated variants. A) Shh, non-palmitoylated C25AShh, and non-lipidated C25AShhN were released into media containing 80 μg/mL HDL; their protein amounts determined by immunoblotting and normalized; and conditioned media added to C3H10T1/2 reporter cells to induce their differentiation into Alp-producing osteoblasts. Mock-treated C3H10T1/2 cells served as non-differentiating controls. At low (1x) and high (2x) concentrations, Shh and C25AShh induced C3H10T1/2 differentiation in a similar manner, as determined by Alp activity measured at 405 nm. Again, C25AShhN was completely inactive, in contrast to bioactive HDL-associated non-palmitoylated C25SShh. One-way ANOVA, Sidaks multiple comparisons test. ****: p<0.0001, ***: p<0.001, n.s. >0.1, n=4. See Table S1 for additional statistical information. B) Similar transcription of Hh target genes Ptch1 and Gli1 by HDL-associated Shh and C25AShh in two independent experiments. C3H10T1/2 reporter cells were stimulated with similar amounts of Shh, C25AShh, and C25AShhN at high (2x) and half (1x) concentrations.

Earlier version: The C-terminal cholesteroylated Shh peptide is required and sufficient for Disp-mediated export and association with HDL

Revised version: The C-terminal cholesteroylated Shh peptide is required and sufficient for Disp-mediated protein export and HDL association

Earlier version: We asked next whether both terminal lipids are required for Disp-mediated Shh export to HDL, or whether the C-cholesteroylated peptide is sufficient for the transfer. The latter possibility was suggested by N-palmitoylated peptide loss during solubilization from Disp-expressing cells in the presence of serum (Figs. 1-3) and by Disp-independent ShhN solubilization into serum-depleted medium (Fig. 1C) 64. We expressed palmitoylated/non-cholesteroylated ShhN and non-palmitoylated/cholesteroylated C25SShh and confirmed that palmitoylated ShhN solubilization from the plasma membrane was independent of Disp and HDL (Fig. 6A, [1] denotes material solubilized into serum-free DMEM, [2] denotes material solubilized by HDL; also see Fig. S7A,B). SEC detected only monomeric ShhN, regardless of the absence or presence of HDL (Fig. 6A’,A’’). These findings suggest that the N-terminal Shh palmitate is not a good substrate for Disp-mediated transfer to HDL. In contrast, only small amounts of monomeric C25SShh were solubilized into serum-free media (Fig. 6B,B’, labeled [3]), and HDL strongly increased soluble C25SShh and its assembly into sizes similar to those of HDL (Fig. 6B’’, labeled [4]; also see Fig. S7C,D). Supporting the importance of Shh cholesteroylation for its transfer to HDL, soluble C25SShh remained lipidated (Fig. S7E). Together, our findings suggest that the C-terminal cholesterol moiety is required for Disp-mediated Shh transfer and HDL association. However, we also noted that a fraction of monolipidated C25SShh solubilized in a Disp independent manner (Fig. 6B, asterisks; Fig. S7F, Fig. 6C, the green dotted line indicates the distribution of C25SShh sizes if released from Disp-/- cells). We explain this finding by unspecific C25SShh release as a consequence of protein overexpression in the presence of HDL and support this claim by a report showing that half-times for spontaneous cholesterol desorption from membranes of several hours are dramatically decreased if large proteins – such as Shh – are attached to a single anchor 3. Once desorped, lipidated proteins either reintegrate into the plasma membrane or associate with lipophilic soluble competotors, as shown in Fig. S7G. This demonstrates that another important role of fully conserved dual N- and C-terminal Hh lipidation is to prevent non-enzymatic morphogen desorption from the plasma membrane.

Revised version: We asked next whether both terminal lipids are required for Disp-mediated Shh export to HDL, or whether the C-cholesteroylated peptide is sufficient for the transfer. The latter possibility was suggested by N-palmitoylated peptide loss during solubilization from Dispexpressing cells in the presence of serum (Figs. 1-3) and by Disp-independent ShhN solubilization into serum-depleted medium (Fig. 1C) 64. We expressed palmitoylated/noncholesteroylated ShhN and non-palmitoylated/cholesteroylated C25SShh and confirmed that palmitoylated ShhN solubilization from the plasma membrane was independent of Disp and HDL (Fig. 7A, [1] denotes material solubilized into serum-free DMEM, [2] denotes material solubilized by HDL; see also Fig. S7A,B). SEC detected only monomeric ShhN, regardless of the absence or presence of HDL (Fig. 7A’,A’’). These findings suggest that the N-terminal Shh palmitate is not a substrate for Disp-mediated transfer to HDL. In contrast, only small amounts of monomeric C25SShh were solubilized into serum-free media (Fig. 7B,B’, labeled [3]), and HDL strongly increased soluble C25SShh and its assembly into sizes similar to those of HDL (Fig. 7B’’, labeled [4]; see also Fig. S7C,D). Supporting the importance of Shh cholesteroylation for its transfer to HDL, soluble C25SShh C-termini remained lipidated (Fig. S7E). Together, our findings suggest that the C-terminal cholesterol moiety is required for Disp-mediated Shh transfer and HDL association. However, we also noted that a fraction of monolipidated C25SShh solubilized in a Disp-independent manner (Fig. 7B, asterisks; Fig. S7F, Fig. 7C; the green dotted line indicates the distribution of C25SShh sizes if released from Disp-/- cells). We explain this finding by unspecific C25SShh release in the presence of HDL, consistent with increased desorption of large proteins – such as Shh – if attached only to a single membrane anchor 3. Once desorbed, lipidated proteins can either reintegrate into the plasma membrane or associate with lipophilic soluble competitors, as shown in Fig. S7G. This demonstrates that another important role of fully conserved dual N- and C-terminal Hh lipidation is to reduce non-enzymatic morphogen desorption from the plasma membrane.

Cholesteroylated C-terminal peptides are required and sufficient for Disp-mediated protein transfer to HDL.

Earlier version: A) Immunoblotted cellular (c) and medium (m) fractions of ShhN-expressing nt ctrl and Disp-/- cells. Shown is unspecific ShhN solubilization into serum-free media (top blot, labeled [1]) or into serum-free DMEM supplemented with HDL (bottom blot, labeled [2]). A’) ShhN [1] expressed under serum-free conditions is solubilized in a monomeric state. A’’) ShhN [2] expressed in the presence of HDL remained monomeric. B) Cholesteroylated C25SShh solubilization into serum-free medium is strongly impaired (top blot, labeled [3]), but increases Disp-dependently in the presence of HDL (bottom blot, labeled [4]). Asterisks denote C25SShh solubilized independent of Disp function. B’) Most C25SShh [3] in serum-free media is monomeric, but a small fraction also assembles into larger units. B’’) C25SShh [4] expressed in the presence of HDL increases in molecular weight to match the molecular weight range of HDL (orange line). C) SEC of C25SShh solubilized from Disp-expressing cells (solid green line) or from Disp-/- cells (dotted line). D) SEC of cholesteroylated mCherry solubilized from nt ctrl cells (solid lines) or from Disp-/- cells (dotted lines) under the same conditions. Dashed lines indicate proteins solubilized under serum-free conditions. Note that most mCherry associates with HDL in a Disp-regulated manner. E) C25SShh (green line) disassembles from HDL in 50% ethanol (bright green line) or in 0.1% Triton X-100 (bright green dashed line). HDL (orange line) disassembly under the same conditions is confirmed by ApoA1 size shift towards the monomeric 32 kDa protein (bright orange line).

Revised version: A) Immunoblotted cellular (c) and medium (m) fractions of ShhN-expressing nt ctrl and Disp-/- cells. Shown is unspecific ShhN solubilization into serum-free media (top blot, labeled [1]) or into serum-free DMEM supplemented with HDL (bottom blot, labeled [2]). A’) ShhN [1] expressed under serum-free conditions is solubilized in a monomeric state. A’’) ShhN [2] expressed in the presence of HDL remained monomeric (e.g. is not HDL associated). B) Cholesteroylated C25SShh solubilization into serum-free medium is strongly impaired (top blot, labeled [3]), but increases Dispdependently in the presence of HDL (bottom blot, labeled [4]). Asterisks denote C25SShh solubilized independent of Disp function. B’) Most C25SShh [3] in serum-free media is monomeric. B’’) C25SShh [4] expressed in the presence of HDL increases in molecular weight to match the molecular weight range of HDL (orange line). C) SEC of C25SShh solubilized from Disp-expressing cells (solid green line) or from Disp-/- cells (dotted line) indicates Disp-independent physical desorption and unregulated HDL association of the monolipidated protein. D) SEC of cholesteroylated mCherry solubilized from nt ctrl cells (solid lines) or from Disp-/- cells (dotted lines) under the same conditions. Dashed lines indicate proteins solubilized under serum-free conditions. Note that most mCherry associates with HDL in a Dispmediated manner. E) C25SShh (green line) disassembles from HDL in 50% ethanol (bright green line) or in 0.1% Triton X-100 (bright green dashed line). HDL (orange line) disassembly under the same conditions is confirmed by ApoA1 size shift toward the monomeric 32 kDa protein (bright orange line).

Earlier version: Is the cholesteroylated peptide sufficient for Disp-mediated transfer to HDL, or does the globular Shh ectodomain participate in the transfer 45, 65? To answer this question, we replaced most of the C25SShh ectodomain with mCherry N-terminally flanked by the Shh secretion signal and C-terminally flanked by the Shh cholesterol transferase/autoprocessing domain. This strategy resulted in the secretion of cholesteroylated mCherry (Fig. S7H) to the cell surface and its association with the outer leaflet of the plasma membrane. As observed previously for Shh (Fig. 5), mCherry remained monomeric or formed small assemblies if released into serum-depleted media (Fig. 6D, dashed line), and HDL shifted mCherry elution into fractions that also contained ApoA1 (solid line). This result supports that HDL serves as a soluble acceptor for C-terminally cholesteroylated peptides. SEC also revealed that loss of Disp function decreased the relative amounts of HDL-associated material (solid line) and increased the relative amounts of solubilized monomeric protein (Fig. 6D, dotted line). Therefore, cholesteroylated peptides represent the critical determinant for Disp-mediated Shh transport to HDL. In contrast, Shh ectodomain interactions with Disp are not essential for the transfer. We supported physical C25SShh interactions with HDL by the addition of EtOH or Triton X-100, both of which disassembled soluble high-molecular weight aggregates into monomers or small assemblies (Fig. 6E).

Revised version: Is the cholesteroylated peptide sufficient for Disp-mediated transfer to HDL, or does the globular Shh ectodomain participate in the transfer 47, 65? To answer this question, we replaced most of the C25SShh ectodomain with mCherry N-terminally flanked by the Shh secretion signal and C-terminally flanked by the Shh cholesterol transferase/autoprocessing domain. This strategy resulted in the secretion of cholesteroylated mCherry to the cell surface and its association with the outer leaflet of the plasma membrane (Fig. S7H). As observed previously for Shh (Fig. 5A), mCherry remained monomeric or formed small assemblies if released into serumdepleted media (Fig. 7D, dashed violet line), and HDL shifted mCherry elution into fractions that also contained ApoA1 (solid violet line). This result supports that HDL serves as a soluble acceptor for C-terminally cholesteroylated peptides. SEC also revealed that loss of Disp function decreased the relative amounts of HDL-associated material (Fig. 7D, solid violet line) and increased the relative amounts of solubilized monomeric protein (Fig. 7D, dotted violet line). Therefore, cholesteroylated peptides represent the critical determinant for Disp-mediated Shh transport to HDL. In contrast, Shh ectodomain interactions with Disp are not essential for the transfer. We supported physical C25SShh interactions with HDL by the addition of EtOH or Triton X-100, both of which disassembled soluble high-MW aggregates into monomers or small assemblies (Fig. 7E, bright green lines indicate C25AShh after treatment and the bright orange line ApoA1 after treatment).

Discussion

Dual Shh lipidation and Disp-regulated Shh solubilization are interdependent

Earlier version: It is firmly established that Hh solubilization and signaling depend on Disp co-expression in source cells. However, the additional involvement of many cell surface and soluble co-factors has complicated efforts to unravel the exact nature of Disp-mediated Hh solubilization. To address this problem, we transfected Disp-deficient and Disp-expressing source cells 4 with Shh and systematically modified Shh solubilization by a number of extracellular co-factors. Then, cell-surface-associated Shh upstream of Disp function and soluble Shh downstream of Disp function were analyzed together with their co-factors by using unbiased biochemical tools. The first important result of this approach is that both Shh lipids act together to prevent uncontrolled Shh solubilization from producing cells. In contrast, artificially produced C25SShh that lacks N-palmitate and ShhN that lacks C-cholesterol are constitutively solubilized – either by non-enzymatic desorption or by unregulated shedding (Fig. 7A). The second important result is that Disp and Scube2 specifically and synergistically increase shedding of the dual-lipidated precursor. These two observations suggest that cell-surface shedding constitutes a “ground state” from which dual-lipidated Shh is protected by tight plasma membrane association of both lipids – but only until Disp and Scube2 render Shh prone to proteolytic processing of the terminal peptides. This concept is important for the reinterpretation of the observation that dual lipidation is essential for unperturbed Hh biofunction in vivo 51, 52, 53, 54, 55. The current interpretation of these observations is that both Hh lipids contribute directly to Ptc receptor signaling for maximal signaling. Our results support an alternative mechanism acting upstream of Ptc: The selective artificial prevention of N- or C-terminal lipidation during Hh/Shh biosynthesis in vivo (achieved by the deletion of the N-terminal palmitate acceptor cysteine or of Hhat, or by prevented C-cholesteroylation 51, 52, 53, 54, 55) may have rendered spatiotemporally controlled Hh/Shh solubilization into unregulated “leakage” of monolipidated precursors from producing cells and tissues (Fig. 1). The “non-availability” of dual-lipidated Shh for Disp-controlled spatio-temporal signaling may then have caused the observed Hh loss-of-function phenotypes 51, 52, 53, 54, 55. Therefore, our release model does not contradict the established essential role of unimpaired dual Shh lipidation during biosynthesis for Hh activity regulation in vivo, but provides an alternative explanation for it. Support for our model of Disp- and Scube2-regulated Hh shedding and release on HDL include the following: Disp is required only in Hh-producing cells and not in receiving cells 64; Hh signaling defects in Disp-deficient model organisms are caused by a defect in the deployment of cell-surface associated Shh 5, 6, 7; and full Disp activity requires Scube2 10. Further support comes from our finding that de-palmitoylated Shh in serum-containing media is fully bioactive (Fig. 4A) and from the in vivo observation that de-palmitoylated Hh variants impair development to variable degrees but can also induce increased ectopic signaling 38, 55.

Revised version: In the past, it was firmly established that Hh solubilization and signaling depend on Disp co-expression in source cells 5, 6, 66. However, the additional involvement of many cell surface and soluble co-factors has complicated efforts to unravel the exact nature of Disp-mediated Hh solubilization. To address this problem, we transfected Disp-deficient and Disp-expressing source cells 4 with Shh and systematically modified Shh solubilization by a number of extracellular co-factors. Cell-surface-associated Shh upstream of Disp function and soluble Shh downstream of Disp function were then analyzed together with their co-factors by using unbiased biochemical tools. The first important result of this approach is that both Shh lipids act together to prevent uncontrolled Shh desorption or shedding from producing cells. In contrast, artificially produced C25S/C25AShh that lacks N-palmitate and ShhN that lacks C-cholesterol are constitutively solubilized either by non-enzymatic desorption or by unregulated shedding (Fig. 8A). The second important result is that Disp and Scube2 specifically and synergistically increase shedding of the dual-lipidated precursor. These two observations suggest that cellsurface shedding constitutes a “ground state” from which dual-lipidated Shh is protected by tight plasma membrane association of both lipids – but only until Disp and Scube2 render Shh prone to proteolytic processing of the terminal lipidated peptides. This concept is important for the reinterpretation of the observation that dual lipidation is essential for unperturbed Hh biofunction in vivo 39, 53, 54, 55, 56. The current interpretation of these observations is that both Hh lipids contribute directly to Ptch1 receptor binding and maximal signaling. Our results support an alternative mechanism acting upstream of Ptch1: We suggest that the selective artificial prevention of N-or C-terminal lipidation during Hh/Shh biosynthesis in vivo (as achieved by the deletion of the N-terminal palmitate acceptor cysteine or of Hhat, or by preventing C-cholesteroylation 39, 53, 54, 55, 56) may have rendered spatiotemporally controlled Hh/Shh solubilization into unregulated “leakage” of monolipidated precursors from producing cells and tissues (Fig. 1). The “non-availability” of dual-lipidated Shh for Disp-controlled spatio-temporal signaling, or the desensitization of target cells to ongoing Shh exposure 67, may then have caused the observed Hh loss-of-function phenotypes 39, 53, 54, 55, 56, 68. For this reason, our release model does not contradict the established essential role of unimpaired dual Shh lipidation during biosynthesis for Hh activity regulation in vivo, but provides an alternative explanation for it. Published support for our model of Disp- and Scube2-regulated Hh shedding and release on HDL includes the following: Disp is required only in Hh-producing cells and not in receiving cells 64; Hh signaling defects in Disp-deficient model organisms are caused by a defect in the deployment of cell-surface associated Shh 5, 6, 7; and full Disp activity requires Scube2 10. Further support comes from the in vivo observation that artificially de-palmitoylated Hh variants impair development to variable degrees, but also show increased signaling in some tissues and induce ectopic signaling 40, 56. The latter observation is difficult to align with proposed essential signaling functions of the Shh palmitate, but is compatible with our model of spatiotemporally perturbed release of monolipidated proteins.

Model of two-way Disp-mediated Shh solubilization.

Earlier version: A) Dual lipidation protects Shh from unregulated cell-surface shedding by tight plasma membrane association of both lipids (blocked shedding is indicated by an x in [1]). In contrast, monolipidated ShhN [2] and C25SShh [3] are prone to unregulated membrane-proximal shedding (indicated by the dashed line) or non-enzymatic desorption. B) Representation of surface hydrophobicity of pdb:7RPH 45 (red) suggests an extended hydrophobic surface channel that may function as a “slide” for lipophiles extending from the plasma membrane (dashed lines) to a cavity of the second extracellular domain (blue arrows). An upward lifted sterol (green stick representation) at the start point of the hydrophobic track may represent an intermediate state of sterol extraction from the membrane, and a lipidic group modeled as the amphiphilic detergent lauryl maltose neopentyl glycol (violet stick structure) may represent the end point of the transfer 45 prior to hand-over to HDL. C) We suggest two sequences of Shh transfer events. In the first event [1], plasma membrane sterol is transferred through the Disp hydrophobic surface channel to HDL acceptors. This process resembles established reverse cholesterol transport. In the second event, if present, C-terminal cholesterol moieties of Shh can also be transferred [2]. This partial Shh extraction exposes the N-terminal cleavage site [3] and makes it prone to proteolytic processing (as is the case for ShhN in A). N-terminal Shh shedding may then release the protein from the plasma membrane [4] to complete the transfer [5]. In addition to, or competing with this process, cholesterol depletion of the plasma membrane (representing the first event, [1]) may trigger shedding of both terminal Shh peptides and the solubilization of monomeric proteins in an indirect manner [6]. See discussion for details.

Revised version: A) Dual lipidation protects Shh from unregulated cell-surface shedding by tight plasma membrane association of both lipids (blocked shedding is indicated by an x in [1]). In contrast, monolipidated ShhN [2] and C25SShh [3] are prone to unregulated membrane-proximal shedding (indicated by the dashed line) or non-enzymatic desorption. B) Representation of surface hydrophobicity of Disp (pdb:7RPH 47) suggests an extended hydrophobic surface channel (hydrophobic residues are shown in red) that may function as a “slide” for lipophiles extending from the plasma membrane (dashed lines) to a cavity of the second extracellular domain (blue arrows). An upward lifted sterol (green stick representation) at the start point of the hydrophobic track may represent an intermediate state of sterol extraction from the membrane, and a lipidic group modeled as the amphiphilic detergent lauryl maltose neopentyl glycol (violet stick structure) may represent the end point of the transfer 47 prior to handover to HDL. C) We suggest two sequences of Shh transfer events. In the first event [1], plasma membrane sterol is transferred through the Disp hydrophobic surface channel to HDL acceptors. This process resembles established reverse cholesterol transport. In the second event, if present, C-terminal cholesterol moieties of Shh can also be transferred [2]. This partial Shh extraction exposes the N-terminal cleavage site [3] and makes it prone to proteolytic processing (similar to ShhN as shown in A). N-terminal Shh shedding may then release the protein from the plasma membrane [4] to complete the transfer [5]. In addition to, or competing with this process, cholesterol depletion of the plasma membrane (representing the first event, [1]) may trigger shedding of both terminal Shh peptides and the solubilization of monomeric proteins in an indirect manner [6]. See Discussion for details.

Disp mediates Shh solubilization by two functionally related transfer modes

Earlier version: The second important result of our study is that the absence of all serum traces reduces or even abolishes Disp/Scube2-mediated Shh solubilization. This reveals that Disp and Scube2 are insufficient to solubilize dual-lipidated Shh and that a previously overlooked serum factor – that we identified as HDL – is also required. We showed previously that Disp knockout in cells impairs 3Hcholesterol efflux and increases cholesterol levels in membranes, and that this in turn inhibits Shh shedding in an indirect manner 4. We also suggested that HDL promotes shedding by serving as a soluble sink for Disp-exported cholesterol 4. Consistent with this possibility, Disp contains a sterol-sensing domain and an extended hydrophobic surface channel that may function as a “slide” for lipophiles. This slide may be powered by a transmembrane Na+ flux 45, similar to the H+ flux that drives related prokaryotic resistance-nodulation-division transporter export of small molecules 66.

Revised version: The second important result of our study is that the absence of all serum traces reduces or even abolishes Disp/Scube2-mediated Shh solubilization. This reveals that Disp and Scube2 are insufficient to solubilize dual-lipidated Shh and that a serum factor is also required. We showed previously that Disp knockout in cells impairs 3Hcholesterol efflux and increases cholesterol levels in membranes, and that this in turn inhibits Shh shedding in an indirect manner 4. We also showed that HDL serves as a soluble sink for Disp-exported cholesterol and that Shh shedding from Disp-/- cells was restored by Disp and the established cholesterol transporter Ptch1 4.

Added in the revised version: Consistent with these published observations, both Disp and Ptch1 contain sterol-sensing domains likely involved in cholesterol transport or in its regulation. Together, these published findings suggest that the required serum factor is HDL.

Added in the revised version: Here we describe that HDL also accepts cholesteroylated – but not dual-lipidated – Shh from Disp (Fig. 8C). This finding is supported by the presence of an extended hydrophobic surface channel in Disp that may function as an open “slide” for lipophiles (Fig. 8B). This slide may be powered by a transmembrane Na+ flux 47, similar to the H+ flux that drives related prokaryotic resistance-nodulation-division transporter export of small molecules 69. Disp exporter function of cholesteroylated proteins to HDL is further supported by the published concept that the fly LPP lipophorin carries cholesteroylated Hh in vivo 13, 14, 15. We extend this published concept by showing that the most C-terminal cholesteroylated Shh peptide is sufficient for direct Disp-mediated protein export to HDL because mCherry linked to this peptide also transfers to HDL in a mostly Disp-dependent manner. This is in line with a previous report demonstrating that C-cholesterol is required and sufficient for Disp-mediated protein export 10.

Earlier version: We describe here that Disp can additionally transfer cholesteroylated Shh to soluble HDL acceptors in a direct manner (Fig. 7C). Both Disp functions confirm and support the established concept that the fly LPP lipophorin carries cholesteroylated Hh in vivo 13, 14, 15, yet that highly bioactive proteins that do not associate with LPPs, that lack the sterol, and that are secreted as dimers and monomers also exist 14. We extend this established concept further by showing that the most C-terminal cholesteroylated Shh peptide is sufficient for direct Disp-mediated protein export to HDL because mCherry linked to this peptide also transfers to HDL in a Disp-dependent manner. This finding is also in line with a previous report demonstrating that C-cholesterol is required and sufficient for Disp-mediated protein export 10. Because of their small size of 5-10 nm, HDLs are not only abundant in the circulation, but are also present in interstitial fluids, both in the adult and throughout development 14. In contrast, larger LPPs such as LDLs and VLDLs are restricted in their distribution by vascular impermeability 50, 67. The ratio of HDL to LDL in interstitial fluid has been estimated to average 50:1 67. These properties make HDL well-suited not only for reverse cholesterol transport but also for the transport of cholesteroylated Hh/Shh cargo to its receptor Ptc – that has also been shown to act as a LPP receptor 68, just like Disp 4. This would lead to the interesting concept that “Hh-free” LPP and “Hh-loaded” LPPs may compete for Ptc receptor binding, the former by promoting Ptc-mediated cholesterol export to suppress signaling and the latter by terminating cholesterol efflux to trigger Hh signaling 41. Indeed, this fascinating mechanism of Ptc activity regulation has already been demonstrated in vitro and in vivo 14.

Revised version: Because of their small size of 5-10 nm, HDLs are not only abundant in the circulation, but are also present in interstitial fluids, both in the adult and throughout development 14. In contrast, larger LPPs such as LDLs and VLDLs are restricted in their distribution by vascular impermeability 52, 70. The ratio of HDL to LDL in interstitial fluids has been estimated to average 50:1 70. These properties make HDL well suited not only for reverse cholesterol transport, but also for the transport of cholesteroylated Hh/Shh cargo to its receptor Ptch1 that also acts as an LPP receptor 71. This would lead to the interesting concept that “Hh-free” LPP and “Hh-loaded” LPPs may compete for Ptch1 receptor binding, the former by promoting Ptch1-mediated cholesterol export to suppress signaling and the latter by terminating cholesterol efflux to trigger Hh signaling 43. Indeed, this elegant mechanism of Ptch1 activity regulation has previously been demonstrated, both in vitro and in vivo 14.

Earlier version: N-terminal shedding during Disp-mediated Hh export is also required in vivo

Revised version: N-terminal shedding during Disp-mediated Hh export in vivo

Earlier version: Our in vitro-derived concept of direct cholesteroylated Hh hand-over from Disp to LPPs and coupled N-terminal peptide shedding is supported by the in vivo finding that site-directed mutagenesis of the sheddase target site fully abolishes mutant Hh biofunction in the developing Drosophila wing and eye (see Fig. S8 for a detailed description of these previously published experiments 34, 35). More importantly, the same mutated proteins suppress endogenous Hh biofunction in a dominant-negative manner 34, 35. We can now explain both observations by mutated Hh transgenes having their C-terminal peptides readily associated with the extended Disp “slide” for lipophiles, but being unable to complete the transfer due to blocked N-terminal shedding and continued plasma membrane association of the Hh N-terminus (preventing step [3] and [4] to freeze [2] in Fig. 7C). As a consequence, the resulting Disp bottleneck would also slow down endogenous Hh release, explaining the observed dominant-negative developmental defects 34, 35 (Fig. S8C,D). This sequence of events is supported by reversed dominant-negative defects of the mutant protein upon additional removal of the N-terminal palmitate membrane anchor 34, 35. The same site-directed mutagenesis approach on the Hh C-terminus, however, affects Hh biofunction only mildly and does not suppress endogenous Hh biofunction in vivo.

Revised version: Our in vitro-derived concept of required N-terminal peptide shedding is supported by the in vivo finding that site-directed mutagenesis of the sheddase target site fully abolishes transgene function in the developing Drosophila wing and eye (see Fig. S8 for a detailed description of repeated and combined experiments from published work 34, 35). Another striking observation was that the same mutated proteins suppress endogenous Hh biofunction in a dominant-negative manner 34, 35. The same site-directed mutagenesis approach on the Hh C-terminus, however, affects biofunction of the transgene only mildly and does not suppress endogenous Hh biofunction in vivo. Both observations can now be explained by mutated Hh transgenes having their C-terminal peptides readily associated with the extended Disp “slide” for lipophiles, with the N-mutated protein unable to complete the transfer due to blocked N-terminal shedding and continued plasma membrane association of the palmitoylated Hh N-terminus (preventing step [3] and [4] because of transfer arrest at step [2] in Fig. 8C). As a consequence, the resulting Disp bottleneck would slow down endogenous Hh release, explaining the observed dominant-negative developmental defects 34, 35 (Fig. S8C,D). This sequence of events is supported by reversed dominant-negative defects of the N-mutant protein upon additional removal of the palmitate membrane anchor 34, 35. Therefore, taken together, our findings suggest that most Hh solubilization in vivo does not require C-terminal Hh shedding, but requires direct cholesteroylated Hh handover from Disp to LPPs. During this process, only the palmitoylated N-terminus is shed to complete the transfer (Fig. 8C, Fig. S8). Our finding that cholesteroylated C-terminal peptides are sufficient for Disp-mediated transfer to HDL is supported by the in vivo observation that transgenes with cholesteroylated 27 kDa green fluorescent protein tags downstream of the 19 kDa Hh signaling domain are bioactive in flies 72 and mice 73. Yet, Disp specificity in vivo can be elegantly explained by Hh being the only known metazoan protein with covalently linked cholesterol. Finally, our finding that palmitoylated Shh N-termini are not extracted and relayed by Disp is supported by the in vivo observation that transgenic expression of ShhN – the variant protein that is N-palmitoylated but lacks C-cholesterol – rescues many early Hh-related embryonic lethal defects in Disp-/- mutant mice 64.

Removed in the revised version: Together, these findings suggest that most Hh solubilization in vivo does not require C-terminal Hh shedding, but requires direct cholesteroylated Hh hand-over from Disp to LPPs. During this process, only the palmitoylated N-terminus is shed to complete the transfer (Fig. 7C, Fig. S8). Our finding that cholesteroylated C-terminal peptides are sufficient for Disp-mediated transfer to HDL is also consistent with the in vivo observation that transgenes with 27 kDa green fluorescent protein (GFP) tags inserted between the 19 kDa Hh signaling domain and its C-terminal cholesteroylated peptide are bioactive in flies 69 and mice 70. We also note that Disp specificity of C-cholesteroylated Hh and Hh-GFP release in vivo can be elegantly explained by Hh being the only known metazoan protein with a covalently linked cholesterol moiety. Finally, our finding that palmitoylated Shh N-termini are not extracted and relayed by Disp (but instead undergo shedding) is supported by the in vivo observation that transgenic expression of ShhN – the variant protein that is N-palmitoylated but lacks C-cholesterol – rescues many early embryonic lethal defects in Disp null mutant mice 64.

Supporting information